[1541214283] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541214368] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.18 ms
[1541214372] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.45 ms
[1541214573] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541215377] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1541215381] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541215405] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1541215458] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541215554] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541215586] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541215591] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.21 ms
[1541215761] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541215861] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541215865] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541215905] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541215952] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541215985] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541215989] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541215998] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.65 ms
[1541216043] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1541216103] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541216193] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541216284] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.29 ms
[1541216410] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541216481] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541216485] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541216680] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1541216755] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541216915] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541217043] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.66 ms
[1541217047] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.88 ms
[1541217131] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541217206] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541217382] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541218341] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541218405] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.30 ms
[1541218453] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.96 ms
[1541218467] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1541218638] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541218644] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541218753] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541218758] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541218898] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541219012] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1541219400] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 232.55 ms
[1541219404] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541219474] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541219490] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1541219494] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541219616] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.46 ms
[1541219635] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541219639] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541219671] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1541219675] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541219768] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1541219778] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.74 ms
[1541219910] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.77 ms
[1541220069] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.77 ms
[1541220210] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541220390] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541220748] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541220811] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541220913] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1541220973] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541220974] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1541221203] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541221262] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.88 ms
[1541221294] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.41 ms
[1541221393] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541221431] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541221472] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1541221712] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541221783] HOST ALERT: sad29-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541221847] HOST ALERT: sad29-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541221892] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541221896] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541222024] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.42 ms
[1541222107] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.57 ms
[1541222293] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1541222479] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541222642] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.38 ms
[1541222652] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.54 ms
[1541222725] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.09 ms
[1541222947] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.41 ms
[1541222989] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1541223041] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1541223495] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.23 ms
[1541223518] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541223559] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.14 ms
[1541223579] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 15.19 ms
[1541223638] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541223690] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.76 ms
[1541223694] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.50 ms
[1541223820] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.74 ms
[1541223824] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541223885] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.35 ms
[1541224010] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.78 ms
[1541224014] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.02 ms
[1541224035] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541224039] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541224097] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541224151] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541224296] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.68 ms
[1541224300] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.80 ms
[1541224428] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.37 ms
[1541224452] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541224522] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.62 ms
[1541224526] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.61 ms
[1541224617] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.82 ms
[1541224728] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.99 ms
[1541224746] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.88 ms
[1541224750] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.45 ms
[1541225208] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1541225253] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1541229244] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541229692] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541229754] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.31 ms
[1541229766] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541231051] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541231115] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1541235709] HOST ALERT: mikro40;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1541235773] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 13.31 ms
[1541237037] HOST ALERT: eliz-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541237101] HOST ALERT: eliz-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541240064] HOST ALERT: mikro44;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541240128] HOST ALERT: mikro44;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541241689] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541241692] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541241812] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541241818] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541242122] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.78 ms
[1541242413] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.05 ms
[1541247105] HOST ALERT: mikro36;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.36.200)
[1541247169] HOST ALERT: mikro36;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.28 ms
[1541247819] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541247835] HOST ALERT: varsh-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541251320] SERVICE ALERT: kirochn-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541251444] SERVICE ALERT: kirochn-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 74.35 ms
[1541254600] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541254724] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541256117] SERVICE ALERT: mikro39;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.200)
[1541256241] SERVICE ALERT: mikro39;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.82 ms
[1541256941] HOST ALERT: lens22-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541257005] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.77 ms
[1541257931] HOST ALERT: lens22-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541257995] HOST FLAPPING ALERT: lens22-k1;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1541257995] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.32 ms
[1541261379] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541261505] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541261754] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541261817] HOST FLAPPING ALERT: lens22-k1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541263620] SERVICE ALERT: nev-k6;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541263650] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541263771] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541263830] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541264699] SERVICE ALERT: sadov-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541264729] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541264850] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541264909] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541265346] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541265472] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541265552] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541266453] HOST ALERT: mikro41;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541266517] HOST ALERT: mikro41;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.76 ms
[1541268217] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541268224] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268227] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541268252] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541268266] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541268343] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541268350] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541268353] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541268367] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541268370] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541268373] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541268378] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541268394] HOST ALERT: blago-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268408] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541268411] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541268416] SERVICE ALERT: lens22-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541268419] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541268426] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268434] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541268437] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541268487] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541268492] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541268494] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541268497] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541268520] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541268527] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541268531] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541268537] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541268539] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541268545] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541268546] SERVICE ALERT: tur-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541268549] HOST ALERT: gavan-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268549] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541268550] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268550] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541268557] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541268558] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541268561] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268563] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541268571] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268606] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541268618] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541268637] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541268662] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541268669] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541268675] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541268675] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541268678] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541268684] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541268687] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541268723] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541268724] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541268730] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541268751] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541268764] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541268793] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541268830] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541268839] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541268839] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541268847] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541268885] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541268904] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541268906] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541268924] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541268965] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541269050] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541269121] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541269174] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541269188] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541269202] SERVICE ALERT: eliz-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541269205] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541269240] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541269247] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541269300] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541269325] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541269328] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541269413] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541269470] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541269521] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541269524] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541269644] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541269647] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541269723] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541269726] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541269744] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541269747] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541269846] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541269852] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541269867] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541269871] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541270092] SERVICE ALERT: mikro57;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541270096] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541270160] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.82 ms
[1541270167] SERVICE ALERT: msk197-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541270170] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541270212] HOST ALERT: mikro42;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.42.200)
[1541270216] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.89 ms
[1541270262] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.98 ms
[1541270290] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541270293] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541271194] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541271320] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541271578] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541272037] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541272040] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541272163] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541272267] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541272270] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541272390] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541272396] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541272414] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541272503] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541272540] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541272629] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541272708] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541272797] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541273834] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541273849] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541273960] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541273965] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541274091] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541274116] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541274234] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541274296] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541274413] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541274642] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541274645] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541274765] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541274768] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541274812] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541274815] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541274930] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541274933] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541274963] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541274996] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541275053] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541275056] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541275131] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275311] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541275330] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541275424] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275456] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541275469] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275480] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541275550] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541275581] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541275581] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541275584] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541275591] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275606] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541275620] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275649] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541275707] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541275710] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541275711] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541275745] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541275771] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541275772] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541275809] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541275880] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541275887] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541275888] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541275889] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541276004] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541286504] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541286568] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.66 ms