[1542067610] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542067635] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542067674] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1542067815] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542067816] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542068000] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542068065] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 10.68 ms
[1542068419] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1542068745] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542068925] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542069742] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 161.23 ms
[1542069746] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1542069890] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542070045] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.34 ms
[1542071272] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542071337] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.04 ms
[1542071558] SERVICE FLAPPING ALERT: mikro56;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1542071567] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542071691] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542072197] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542072295] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542072470] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542072650] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542072785] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542072918] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542072921] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542072955] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 2.07 ms
[1542073069] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542073101] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542073280] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542073345] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.84 ms
[1542073509] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542073642] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1542073642] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1542073646] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1542073970] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542073972] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542074097] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1542074099] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.20 ms
[1542074136] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542074261] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.51 ms
[1542074466] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542074646] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542074725] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542074864] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.83 ms
[1542074869] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542074905] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542074993] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1542074997] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1542075323] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542075467] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542075503] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542075531] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1542075598] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 5.24 ms
[1542075857] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542076037] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542076087] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542076211] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1542076211] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542076215] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1542076220] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542076223] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542076371] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542076378] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.10 ms
[1542077003] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542077141] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542077183] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542077322] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542077415] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1542077544] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1542077548] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542077712] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542077776] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.46 ms
[1542077874] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542078026] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1542078149] HOST ALERT: dek-k1;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 2.26 ms
[1542078402] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542078582] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542078610] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542078735] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 11.25 ms
[1542078771] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542078774] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542078895] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1542078899] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1542079039] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1542079225] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542079361] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542079365] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542079405] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542079485] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1542079489] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1542079499] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542079684] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542079712] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1542079815] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542079825] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542079842] SERVICE ALERT: nev6-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.99 ms
[1542079846] HOST ALERT: nev6-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.28 ms
[1542079914] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542079920] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1542079945] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1542079956] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542079963] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1542079995] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542080005] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542080089] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1542080099] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542080138] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.57 ms
[1542080244] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1542080248] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.98 ms
[1542080250] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1542080265] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1542080270] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1542080294] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1542080295] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1542080321] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.74 ms
[1542080325] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.87 ms
[1542080353] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1542080372] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.95 ms
[1542080375] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1542080376] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.66 ms
[1542080415] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542080423] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1542080427] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.36 ms
[1542080494] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1542080501] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.83 ms
[1542080522] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.73 ms
[1542080544] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1542080595] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542080603] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1542080647] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1542080651] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1542080689] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1542080725] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542080763] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1542080767] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1542080785] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1542080790] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 6.83 ms
[1542080813] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.81 ms
[1542080852] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1542080921] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1542081037] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.26 ms
[1542081119] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 18.27 ms
[1542081233] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1542081276] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1542081322] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542081473] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542081597] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542081652] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542081832] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542081842] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.14 ms
[1542081846] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.09 ms
[1542081886] SERVICE ALERT: mikro56;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.32 ms
[1542081889] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542081953] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1542082010] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542082063] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.90 ms
[1542082136] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1542082194] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 1.93 ms
[1542082279] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542082344] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.22 ms
[1542082970] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542083035] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.84 ms
[1542083062] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542083127] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 7.79 ms
[1542083476] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1542083480] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1542083540] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.58 ms
[1542083647] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1542083722] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1542083779] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1542083815] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.77 ms
[1542083898] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.52 ms
[1542083931] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542083995] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542084013] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1542084025] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542084045] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.66 ms
[1542084080] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1542084089] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 16.74 ms
[1542084108] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1542084298] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542084362] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542084415] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542084595] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542084601] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542084663] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542084689] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1542084704] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.42 ms
[1542084727] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542084768] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1542084813] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542084817] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542084819] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.56 ms
[1542084823] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.66 ms
[1542084927] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542085087] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542085107] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542085201] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1542085353] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542085533] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.22 ms
[1542085533] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542085637] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542085703] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.11 ms
[1542085761] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542085765] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542085824] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 11.01 ms
[1542085824] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542085854] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542085875] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.41 ms
[1542085919] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 8.12 ms
[1542085948] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1542086183] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.82 ms
[1542086214] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.15 ms
[1542086244] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542086388] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542086391] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542086424] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542086459] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.21 ms
[1542086508] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.30 ms
[1542086513] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.17 ms
[1542086517] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542086575] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542086699] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.10 ms
[1542086703] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.80 ms
[1542086843] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542086908] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.33 ms
[1542087029] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542087113] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1542087144] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.69 ms
[1542087148] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1542087152] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.77 ms
[1542087161] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1542087184] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 6.81 ms
[1542087404] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.80 ms
[1542087414] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1542087440] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.25 ms
[1542087453] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.14 ms
[1542087456] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.48 ms
[1542087465] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1542087519] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1542087616] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1542087739] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542087740] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542087740] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1542087777] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1542087806] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1542087863] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542087905] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.67 ms
[1542087909] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.73 ms
[1542087920] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542087926] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542087956] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542087969] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1542087973] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.57 ms
[1542087977] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.53 ms
[1542088020] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.34 ms
[1542088042] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542088050] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.27 ms
[1542088089] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.45 ms
[1542088096] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542088106] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542088114] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.60 ms
[1542088233] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1542088301] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1542088303] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.21 ms
[1542088308] HOST ALERT: dek-k2;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 2.27 ms
[1542088346] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542088399] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.18 ms
[1542088404] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542088460] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542088468] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1542088473] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542088585] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.57 ms
[1542088595] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.08 ms
[1542088650] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542088775] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.50 ms
[1542088863] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.32 ms
[1542088953] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 17.30 ms
[1542088957] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.38 ms
[1542089066] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542089130] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542089233] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (21.8% change > 20.0% threshold)
[1542089233] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542089280] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542089378] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542089413] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1542089507] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542089704] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542089833] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542089884] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542089978] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1542089988] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 4.42 ms
[1542090304] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542090383] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542090386] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542090484] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542090508] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 3.23 ms
[1542090512] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542090579] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 60%, RTA = 2.33 ms
[1542090613] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542090793] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542090838] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542090903] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.19 ms
[1542091307] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1542091528] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542091593] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.41 ms
[1542091606] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542091670] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542091996] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542092151] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 3.72 ms
[1542092334] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542092458] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1542092704] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542092831] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1542092897] SERVICE ALERT: mikro59;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542092900] HOST ALERT: mikro59;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.200)
[1542092916] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542093020] SERVICE ALERT: mikro59;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.200)
[1542093026] HOST ALERT: mikro59;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.200)
[1542093042] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1542093076] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542093107] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542093157] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542093212] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542093337] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542093433] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542093578] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542093613] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542093707] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542093837] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542093901] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542093904] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542094084] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.02 seconds)
[1542094179] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.24 ms
[1542094227] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542094285] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542094292] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.47 ms
[1542094330] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542094333] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542094409] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1542094455] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.82 ms
[1542094459] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.23 ms
[1542097479] HOST FLAPPING ALERT: dek-k1;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1542097799] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1542097863] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1542097891] HOST FLAPPING ALERT: mikro56;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1542098054] HOST FLAPPING ALERT: dek-k2;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1542098374] SERVICE FLAPPING ALERT: dek-k1;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1542105021] SERVICE ALERT: mikro59;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542105025] HOST ALERT: mikro59;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542105046] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1542105336] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542105809] SERVICE FLAPPING ALERT: mikro56;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1542108101] SERVICE ALERT: mikro72;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 5.84 ms
[1542108225] SERVICE ALERT: mikro72;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1542108626] HOST ALERT: mikro59;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.200)
[1542108648] SERVICE ALERT: mikro59;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542108672] HOST ALERT: suvor42-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542108737] HOST ALERT: suvor42-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.48 ms
[1542108780] HOST ALERT: mikro59;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.200)
[1542109084] HOST ALERT: mikro59;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542109372] SERVICE ALERT: mikro59;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1542117897] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542117961] HOST ALERT: kamen-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1542118207] HOST ALERT: lens22-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542118271] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1542120612] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542120738] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1542120904] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542121282] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1542121342] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.90 ms
[1542121346] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1542121628] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.57 ms
[1542124619] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1542124745] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1542124913] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1542126356] SERVICE ALERT: nev6-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1542126359] HOST ALERT: nev6-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1542126479] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1542126485] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1542127745] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1542127871] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1542128041] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1542128930] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542129056] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1542129086] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542129564] SERVICE ALERT: turist28-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 3.97 ms
[1542129573] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542129672] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542129688] SERVICE ALERT: turist28-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1542129754] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542132209] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132262] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542132265] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542132267] HOST ALERT: blago-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132310] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542132335] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1542132370] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1542132388] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1542132391] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1542132393] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1542132427] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542132430] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542132478] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542132480] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542132496] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1542132512] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542132529] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1542132550] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542132553] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1542132562] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132570] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542132587] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542132589] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132612] HOST ALERT: lunach-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132623] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542132632] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1542132638] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1542132658] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542132660] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542132685] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1542132693] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1542132696] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1542132707] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542132713] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1542132735] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542132738] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1542132744] HOST ALERT: garkav-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.21)
[1542132747] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542132762] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542132764] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1542132795] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542132798] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542132860] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542132874] HOST ALERT: eliz-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132887] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542132888] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1542132916] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542132918] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542132924] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1542133000] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1542133013] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1542133060] SERVICE ALERT: gavan-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542133063] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542133066] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1542133081] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542133144] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542133150] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542133183] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542133186] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1542133187] SERVICE ALERT: pion12-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.53.21)
[1542133190] HOST ALERT: pion12-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.53.21)
[1542133192] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1542133258] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542133261] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542133270] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1542133276] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1542133295] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542133303] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542133311] SERVICE ALERT: pion12-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542133315] HOST ALERT: pion12-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1542133339] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1542133340] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1542133352] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542133355] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542133425] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542133475] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542133481] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1542133502] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542133628] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1542133771] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542133934] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542133976] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542134058] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1542134100] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1542134186] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542134310] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1542134337] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542134450] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542134453] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542134463] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1542134474] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542134554] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542134573] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542134579] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1542134657] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542134699] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542134783] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1542135117] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542135143] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542135243] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1542135274] SERVICE ALERT: nev-k11;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542135277] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542135340] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542135399] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542135403] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542135425] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542135457] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542135464] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542135520] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542135523] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542135529] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1542135975] SERVICE ALERT: sad29-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542135978] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542136098] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542136101] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1542136132] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542136212] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542136258] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1542136336] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542136462] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1542136625] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542138126] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542138252] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1542138568] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1542138588] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542138590] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542138694] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1542138926] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1542138929] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1542139030] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542139049] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1542139051] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542139052] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1542139210] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542139318] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542139321] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542139325] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542139338] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542139390] SERVICE ALERT: kamen-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542139393] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542139441] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542139444] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1542139462] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1542139505] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542139513] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542139519] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1542139582] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542139608] HOST ALERT: nev6-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1542139621] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542139708] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1542139734] HOST ALERT: nev6-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.22)
[1542139747] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1542139787] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542139802] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542139841] SERVICE ALERT: nev6-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.22)
[1542139878] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542139955] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542140358] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542140443] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542140473] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542140515] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542140538] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542140594] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542140653] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542140660] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542140695] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542140703] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542153550] SERVICE ALERT: kirochn-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.4.22)
[1542153582] SERVICE ALERT: mikro04;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.4.200)