[1541900076] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541900292] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541900356] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541900402] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541900492] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541900559] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541900863] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541901040] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541901161] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1541901161] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541901191] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541901225] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541901315] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541901370] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541901497] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1541901497] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541901551] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541901701] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.28 ms
[1541901823] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541901888] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 29.83 ms
[1541901997] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541902123] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541902295] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541905418] HOST FLAPPING ALERT: mikro56;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541905515] HOST FLAPPING ALERT: dek-k1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541905805] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541905911] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541905924] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.77 ms
[1541906075] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1541907696] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 180.98 ms
[1541907788] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541907809] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541907820] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 35.70 ms
[1541907967] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1541907980] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541908023] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541908027] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541908068] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1541908070] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1541908073] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.59 ms
[1541908074] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541908086] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1541908087] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541908110] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541908119] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.53 ms
[1541908224] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1541908231] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.63 ms
[1541908327] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541908336] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541908383] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541908426] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.73 ms
[1541908470] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1541908473] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541908477] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541908498] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541908525] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.85 ms
[1541908555] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1541908655] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.74 ms
[1541908694] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541908699] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.52 ms
[1541908700] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541908717] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541908745] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541908942] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541909012] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.23 ms
[1541909706] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.08 ms
[1541909782] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541909996] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.32 ms
[1541910347] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1541910778] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541910808] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541910959] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541911063] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 214.96 ms
[1541911067] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.00 ms
[1541911124] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1541911230] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541911250] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541911789] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1541911876] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541912083] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1541912143] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541912147] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541912167] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541912174] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541912178] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541912205] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541912209] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.44 ms
[1541912248] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.65 ms
[1541912569] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.73 ms
[1541912583] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.48 ms
[1541912820] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541912824] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541912850] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.30 ms
[1541912892] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.75 ms
[1541913151] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.09 ms
[1541913315] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.85 ms
[1541913319] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541913372] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.70 ms
[1541913376] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.65 ms
[1541913427] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541913434] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.38 ms
[1541913539] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1541913697] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.94 ms
[1541913718] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1541913984] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.98 ms
[1541914301] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1541914331] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.64 ms
[1541914350] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541914354] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541914600] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541914608] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541914617] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.71 ms
[1541914629] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 13.52 ms
[1541914674] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541914707] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.90 ms
[1541914711] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.73 ms
[1541914894] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.73 ms
[1541915123] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.47 ms
[1541915194] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1541915196] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1541915222] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541915229] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.63 ms
[1541915253] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.03 ms
[1541915331] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.32 ms
[1541915357] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541915367] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1541915444] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541915490] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541915509] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.54 ms
[1541915523] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541915650] HOST ALERT: eliz-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541915956] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541916250] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541917680] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541917683] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541917803] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541917806] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541917835] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.56 ms
[1541917839] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.46 ms
[1541918404] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.17 ms
[1541918408] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.50 ms
[1541918707] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541918833] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541919003] SERVICE FLAPPING ALERT: aviakonstr-k2;Ping;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1541919003] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541919127] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.61 ms
[1541919131] HOST FLAPPING ALERT: aviakonstr-k2;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1541919131] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.56 ms
[1541922131] HOST FLAPPING ALERT: aviakonstr-k2;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541926829] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541926893] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.12 ms
[1541928092] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541928145] HOST FLAPPING ALERT: mikro72;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1541928145] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1541930527] SERVICE FLAPPING ALERT: aviakonstr-k2;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541931745] HOST FLAPPING ALERT: mikro72;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1541934450] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541934453] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541934513] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541934517] HOST ALERT: nev-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541934574] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541934693] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541934763] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541934887] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541934891] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541939794] SERVICE ALERT: mikro59;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.59.200)
[1541939918] SERVICE ALERT: mikro59;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541939940] SERVICE ALERT: bryanc20-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 6.45 ms
[1541940064] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541946171] HOST ALERT: mikro72;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541946235] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.89 ms
[1541949382] SERVICE ALERT: mikro69;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541949506] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.02 ms
[1541949545] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541949609] HOST ALERT: nevsky130-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1541951608] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541951734] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541951867] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541951934] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541951998] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.10 ms
[1541953195] SERVICE ALERT: sad29-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541953198] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541953318] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541953324] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541953607] SERVICE ALERT: nev-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541953637] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541953758] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541953817] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541954331] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1541954336] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541955838] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541955868] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541955989] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541956048] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541957082] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541957085] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541957205] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541957208] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541959393] HOST ALERT: lunach-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541959449] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541959452] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541959483] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541959518] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541959549] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541959572] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541959578] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541959609] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541959612] SERVICE ALERT: engels-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541959621] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541959623] HOST ALERT: engels-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541959630] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541959647] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541959661] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541959677] SERVICE ALERT: engels51-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541959680] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541959706] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541959709] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541959730] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1541959735] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541959743] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541959747] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541959749] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541959756] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541959766] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541959773] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541959800] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541959803] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541959829] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541959833] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541959843] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541959863] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541959869] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541959891] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541959900] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541959904] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541959908] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541959926] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541959949] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541959952] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541959968] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541960015] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541960026] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541960072] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541960078] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541960084] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541960138] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541960201] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541960216] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541960232] SERVICE ALERT: sadov-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541960235] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541960264] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541960301] SERVICE ALERT: varsh-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541960304] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541960309] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541960319] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541960322] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541960330] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541960372] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541960383] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541960415] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541960421] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541960424] HOST ALERT: blago-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541960428] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541960434] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541960442] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541960445] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541960475] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541960489] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541960489] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541960492] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541960550] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541960573] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541960578] SERVICE ALERT: gavan-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541960581] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541960601] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541960612] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541960618] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541960637] HOST ALERT: nevsk104-k4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541960649] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541960701] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541960707] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541960754] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541960772] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541960775] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541960819] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541960897] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541960900] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541961020] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541961023] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541961157] HOST ALERT: bkuna-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541961283] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541961423] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541961428] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541961549] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541961874] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541962723] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541962734] HOST ALERT: kamen-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541962846] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541962860] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541963146] SERVICE ALERT: sad29-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541963149] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541963269] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541963275] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541963517] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541963520] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541963640] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541963643] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541963714] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541963714] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541963717] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541963717] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541963766] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541963769] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541963837] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541963843] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541963865] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541963889] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541963895] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541963897] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541964300] SERVICE ALERT: nev-k7;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541964303] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541964451] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541964483] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541964754] SERVICE ALERT: nev-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541964784] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541964861] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541964905] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541964964] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541965041] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541965158] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541965823] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541965826] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541965946] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541965949] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541966004] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541966124] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541966127] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541966130] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541966247] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541966250] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541966429] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541966538] SERVICE ALERT: msk197-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541966541] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541966645] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541966661] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541966667] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541966758] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541966771] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541966793] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541966854] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541966872] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541966884] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541966909] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541966912] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541966919] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541966943] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541967001] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541967032] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541967038] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541967049] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541967123] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541979819] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1541979823] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541980213] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541980393] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541980418] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541980421] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541980485] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541980515] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541980542] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541980639] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.81 ms
[1541980643] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 28.37 ms