[1542153674] SERVICE ALERT: kirochn-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 64.62 ms
[1542153706] SERVICE ALERT: mikro04;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 56.59 ms
[1542165837] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542165879] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542166081] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 8.55 ms
[1542166312] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1542166366] SERVICE ALERT: nev6-k2;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 1.93 ms
[1542166374] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1542166402] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.54 ms
[1542166478] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.02 ms
[1542166514] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1542166775] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.57 ms
[1542166788] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.25 ms
[1542166926] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1542167042] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1542167067] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.87 ms
[1542167184] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542167227] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1542167250] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1542167495] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1542167511] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1542167514] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1542167567] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1542169179] SERVICE ALERT: mikro68;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.68.200)
[1542169253] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.75 ms
[1542169302] SERVICE ALERT: mikro68;Ping;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (192.168.68.200)
[1542169396] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1542169425] SERVICE ALERT: mikro68;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (192.168.68.200)
[1542169427] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.52 ms
[1542169431] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 10.22 ms
[1542169460] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1542170029] SERVICE ALERT: mikro68;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1542170165] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1542170299] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.28 ms
[1542170536] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542170611] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1542170840] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.82 ms
[1542171219] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542171349] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542171798] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.24 ms
[1542172184] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 18.47 ms
[1542172723] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.44 ms
[1542172749] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1542173088] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.04 ms
[1542173119] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.31 ms
[1542173351] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.58 ms
[1542173547] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.51 ms
[1542173663] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.28 ms
[1542173664] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.77 ms
[1542173671] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.32 ms
[1542173784] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542173823] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.30 ms
[1542173834] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.19 ms
[1542174004] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1542174157] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542174182] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.07 ms
[1542174276] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.91 ms
[1542174817] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1542175384] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.52 ms
[1542175598] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542182408] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.83 ms
[1542182418] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.48 ms
[1542182423] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 21.76 ms
[1542182424] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 28.76 ms
[1542182432] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1542182446] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.09 ms
[1542182447] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1542182453] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.79 ms
[1542182457] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.74 ms
[1542182469] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1542182469] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542182479] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1542182485] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1542182488] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542182488] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1542182492] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.08 ms
[1542182499] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1542182502] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1542182506] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.13 ms
[1542182511] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.20 ms
[1542182512] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 18.40 ms
[1542182516] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.39 ms
[1542182527] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.00 ms
[1542182528] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1542182531] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1542182537] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.85 ms
[1542182538] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1542182540] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.82 ms
[1542182544] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1542182551] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.21 ms
[1542182554] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1542182561] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542182561] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.08 ms
[1542182565] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.77 ms
[1542182569] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542182569] HOST ALERT: nev6-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1542182571] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1542182571] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1542182578] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.96 ms
[1542182584] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1542182588] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1542182594] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1542182600] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.44 ms
[1542182602] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 16.24 ms
[1542182603] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.97 ms
[1542182607] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.49 ms
[1542182621] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.19 ms
[1542182621] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.99 ms
[1542182627] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1542182653] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.82 ms
[1542182655] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.63 ms
[1542182679] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1542182688] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.93 ms
[1542182692] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1542182976] HOST ALERT: mikro04;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542183020] HOST ALERT: kirochn-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542183041] HOST ALERT: mikro04;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 68.58 ms
[1542183085] HOST ALERT: kirochn-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 60.73 ms
[1542183693] HOST ALERT: kirochn-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542183701] SERVICE ALERT: kirochn-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542183710] HOST ALERT: kirochn-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542183717] SERVICE ALERT: kirochn-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.21)
[1542183733] SERVICE ALERT: mikro04;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542183736] HOST ALERT: mikro04;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.200)
[1542183819] HOST ALERT: kirochn-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.22)
[1542183836] HOST ALERT: kirochn-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.21)
[1542183856] SERVICE ALERT: mikro04;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.200)
[1542183862] HOST ALERT: mikro04;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.200)
[1542184125] HOST ALERT: kirochn-k2;UP;HARD;1;PING OK - Packet loss = 60%, RTA = 63.74 ms
[1542184167] HOST FLAPPING ALERT: mikro04;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1542184167] HOST ALERT: mikro04;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 87.69 ms
[1542184426] SERVICE ALERT: kirochn-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 70.25 ms
[1542184441] HOST FLAPPING ALERT: kirochn-k1;STARTED; Host appears to have started flapping (21.7% change > 20.0% threshold)
[1542184441] HOST ALERT: kirochn-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 61.92 ms
[1542184442] SERVICE ALERT: kirochn-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 73.05 ms
[1542184457] SERVICE ALERT: mikro04;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 69.36 ms
[1542187140] HOST FLAPPING ALERT: kirochn-k1;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1542187457] HOST FLAPPING ALERT: mikro04;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1542190395] HOST ALERT: 8line-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542190414] HOST ALERT: 8line-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542201497] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542201561] HOST ALERT: nev-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1542202327] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1542202391] HOST ALERT: mikro69;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 9.46 ms
[1542210714] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542210894] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542211011] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542211498] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1542211632] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542211735] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1542211785] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1542211824] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542212004] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542212036] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1542212362] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542212495] HOST ALERT: sad29-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542212621] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1542212764] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1542213132] SERVICE ALERT: mikro39;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542213256] SERVICE ALERT: mikro39;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1542214215] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542214218] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542214366] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542214398] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542214628] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542214692] HOST ALERT: nev-k11;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542215913] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542216039] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1542216209] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542218606] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542218611] SERVICE ALERT: garkav-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542218614] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542218628] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542218654] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542218675] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542218678] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542218687] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542218707] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542218715] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542218732] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1542218762] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542218780] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1542218783] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542218786] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542218794] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542218798] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542218804] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1542218813] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1542218822] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542218825] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542218833] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1542218833] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542218836] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542218839] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.57 ms
[1542218849] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542218860] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1542218906] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542218912] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1542218945] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542218950] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542218951] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1542218956] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542218962] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1542218975] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1542218983] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542218986] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1542219003] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542219017] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1542219091] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542219143] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1542219146] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542219157] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1542219216] SERVICE ALERT: lens22-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1542219217] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1542219218] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542219219] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1542219220] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542219221] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542219270] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542219287] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542219339] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1542219341] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542219344] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1542219345] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1542219348] SERVICE ALERT: eliz-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542219351] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542219400] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542219438] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1542219442] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542219447] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542219467] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542219471] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542219477] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1542219477] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542219566] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1542219573] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1542219603] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1542219643] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1542219680] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542219683] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542219711] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542219711] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542219769] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1542219773] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542219781] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542219803] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1542219806] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1542219837] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1542219854] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542219862] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1542219863] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1542219891] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542219917] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542219920] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542220040] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542220046] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1542220118] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542220298] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542220625] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542220660] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542220784] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1542220997] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542221000] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542221077] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542221120] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542221123] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1542221297] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542221423] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1542221666] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542222202] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542222252] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542222328] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1542222378] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1542222439] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542222498] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542222522] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542222525] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542222548] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542222565] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1542222645] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542222648] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1542222735] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542222780] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542222906] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1542223049] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1542223372] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542223430] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542223552] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542223942] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542223945] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542224065] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542224071] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1542224810] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1542224936] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1542225111] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1542225705] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542225736] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542225738] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542225741] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542225753] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542225799] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542225818] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542225841] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542225844] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542225861] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542225862] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1542225867] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1542225876] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542225885] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542225942] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1542225964] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542225967] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1542226051] SERVICE ALERT: nev-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542226081] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542226202] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542226203] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542226205] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542226217] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542226261] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542226294] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542226297] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542226343] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1542226354] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542226385] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542226417] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542226420] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1542226513] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542226994] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542227024] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542227145] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542227204] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542227557] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542227737] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542227854] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out