[1541388417] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1541388469] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541388479] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541388483] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1541388510] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541388534] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1541388602] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541388636] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541388656] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541388688] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541388693] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.81 ms
[1541388696] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.88 ms
[1541388712] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541388889] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541388893] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1541388927] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541388929] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541388933] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1541388938] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541388963] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541388967] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1541389037] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541389041] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541389059] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1541389087] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541389091] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541389150] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.25 ms
[1541389166] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 31.84 ms
[1541389177] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.89 ms
[1541389181] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1541389228] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541389296] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1541389386] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541389419] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541389423] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.75 ms
[1541389456] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.63 ms
[1541389527] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1541389605] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541389681] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541389685] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541389814] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541390123] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.00 ms
[1541390137] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541390414] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541390599] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.92 ms
[1541390603] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541390793] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541391087] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.95 ms
[1541391219] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.12 ms
[1541391285] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541391337] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.03 ms
[1541391514] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541391522] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.66 ms
[1541391575] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541392510] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.83 ms
[1541392514] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.74 ms
[1541393034] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541393045] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541393049] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541393306] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1541393317] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541393490] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1541393570] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.30 ms
[1541393800] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541394035] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541394325] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1541394817] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.28 ms
[1541395111] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.69 ms
[1541395113] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.19 ms
[1541395132] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.23 ms
[1541395425] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.56 ms
[1541395429] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.12 ms
[1541395514] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.94 ms
[1541395518] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541395755] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.75 ms
[1541395848] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1541395955] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.72 ms
[1541396087] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.46 ms
[1541396088] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.73 ms
[1541396112] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1541396178] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541396305] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541396309] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541396345] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.98 ms
[1541396463] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541396534] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.47 ms
[1541396585] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 313.57 ms
[1541396591] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.11 ms
[1541396593] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.12 ms
[1541396595] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.50 ms
[1541396597] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.56 ms
[1541396597] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1541396635] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.82 ms
[1541396706] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.96 ms
[1541396716] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.01 ms
[1541396790] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541396798] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.41 ms
[1541396811] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541396815] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1541396884] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1541396969] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.43 ms
[1541397010] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.97 ms
[1541397054] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.99 ms
[1541397058] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.68 ms
[1541397081] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1541397174] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.87 ms
[1541397671] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541397976] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541398408] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.52 ms
[1541398753] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541399166] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1541399269] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541406115] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541406145] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541406209] HOST ALERT: nev-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.26 ms
[1541406239] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1541407264] SERVICE ALERT: mikro70;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541407267] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541407387] SERVICE ALERT: mikro70;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541407390] HOST ALERT: mikro70;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.200)
[1541407408] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541407534] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541407704] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541408142] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541408322] HOST ALERT: prosv48-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541408414] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.75 ms
[1541409313] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541409439] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541409609] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541409977] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541410041] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541414428] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541414432] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1541414588] SERVICE ALERT: mikro70;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541414592] HOST ALERT: mikro70;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541415733] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.22 ms
[1541415737] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.70 ms
[1541416405] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541416469] HOST ALERT: eliz-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541424188] SERVICE ALERT: mikro70;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541424312] SERVICE ALERT: mikro70;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.07 ms
[1541430984] HOST ALERT: mikro60;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.60.200)
[1541431048] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541433533] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541433599] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541433659] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541436410] HOST ALERT: mikro46;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.200)
[1541436449] HOST ALERT: mikro60;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.60.200)
[1541436474] HOST ALERT: mikro46;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.88 ms
[1541436502] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541436506] SERVICE ALERT: nev-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541436509] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541436513] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1541436628] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541436657] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541436689] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541436798] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541436813] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541436843] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541436964] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541437023] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541437036] SERVICE ALERT: nev6-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541437039] HOST ALERT: nev6-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541437159] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541437165] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1541437674] HOST ALERT: mikro46;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.200)
[1541437739] HOST FLAPPING ALERT: mikro46;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1541437739] HOST ALERT: mikro46;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541439837] SERVICE ALERT: mikro61;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541439961] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541440868] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541440982] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541440994] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1541441001] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541441068] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541441105] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541441108] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541441109] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541441112] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541441175] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541441194] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541441216] HOST ALERT: suvor42-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441225] SERVICE ALERT: garkav-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541441231] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541441232] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541441239] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541441255] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441282] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541441294] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541441307] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541441311] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541441311] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541441316] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541441328] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541441331] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541441341] SERVICE ALERT: sadov-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541441342] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541441345] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441348] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541441348] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541441352] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541441368] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541441408] HOST ALERT: garkav-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.21)
[1541441408] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541441412] HOST ALERT: 8line-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1541441420] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541441433] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541441434] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541441437] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541441442] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541441447] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541441454] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541441455] SERVICE ALERT: varsh-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541441458] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441463] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441481] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541441492] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541441505] HOST FLAPPING ALERT: mikro46;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541441514] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541441516] HOST ALERT: eliz-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441525] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541441560] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541441569] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541441573] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541441578] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541441581] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541441590] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541441596] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541441624] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541441640] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541441640] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541441644] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541441711] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541441711] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541441724] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541441777] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541441799] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541441810] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541441836] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541441891] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.29 seconds)
[1541441901] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541441949] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541441961] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541441987] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541442002] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541442011] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541442062] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541442113] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541442130] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541442137] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541442192] SERVICE ALERT: tur-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541442195] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541442234] SERVICE ALERT: engels-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541442237] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541442244] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541442286] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541442315] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541442318] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541442357] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541442360] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541442957] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541443083] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541443253] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541443526] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541443652] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541443811] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541444149] HOST ALERT: kamen-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541444275] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541444413] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541444713] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541444716] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541444792] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541444792] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541444836] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541444842] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541444916] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541444918] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541444918] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541444919] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541444954] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541445039] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541445045] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541445080] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541445086] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541445096] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541445154] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541445235] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541445266] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541445415] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541445496] SERVICE ALERT: nev-k11;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541445526] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541445619] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541445647] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541445706] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541445745] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541445888] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541446654] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541446718] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.33 ms
[1541446859] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541446982] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541446984] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541447110] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541447128] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541447190] SERVICE ALERT: nev-k7;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541447193] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541447280] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541447341] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541447373] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541447678] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541447804] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541447975] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541448131] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541448218] SERVICE ALERT: msk197-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541448221] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541448257] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541448326] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541448329] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541448336] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541448341] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541448347] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541448443] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541448449] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541448455] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541448495] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541448565] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541448623] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541448627] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541448630] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541448675] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541448683] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541448691] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541448723] SERVICE ALERT: nev-k6;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541448739] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541448750] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541448753] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541448756] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541448855] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541448874] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541448933] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541452738] HOST ALERT: mikro40;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1541452802] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1541455561] SERVICE ALERT: mikro61;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541455685] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.97 ms