[1541635273] SERVICE FLAPPING ALERT: mikro52;Ping;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1541635273] SERVICE ALERT: mikro52;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541635389] SERVICE FLAPPING ALERT: mikro69;Ping;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1541635389] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.67 ms
[1541636564] HOST ALERT: mikro04;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 66.74 ms
[1541636584] HOST ALERT: kirochn-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 61.54 ms
[1541636593] HOST ALERT: kirochn-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 43.32 ms
[1541636594] SERVICE ALERT: mikro04;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 46.62 ms
[1541636890] SERVICE ALERT: kirochn-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 67.99 ms
[1541637098] SERVICE ALERT: kirochn-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 53.65 ms
[1541637735] HOST FLAPPING ALERT: mikro61;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541637789] HOST FLAPPING ALERT: mikro69;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1541637954] HOST FLAPPING ALERT: mikro52;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541638035] HOST FLAPPING ALERT: mikro57;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1541638099] HOST FLAPPING ALERT: dach-k1;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1541638189] HOST FLAPPING ALERT: mikro48;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1541641188] SERVICE ALERT: mikro70;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541641312] SERVICE ALERT: mikro70;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541644383] HOST ALERT: kirochn-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.22)
[1541644390] HOST ALERT: mikro04;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541644393] SERVICE ALERT: mikro04;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.200)
[1541644418] HOST ALERT: kirochn-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541644509] HOST ALERT: kirochn-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.22)
[1541644516] HOST ALERT: mikro04;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.200)
[1541644544] HOST ALERT: kirochn-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.21)
[1541644694] HOST ALERT: kirochn-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 58.18 ms
[1541644813] HOST ALERT: kirochn-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 75.93 ms
[1541644820] HOST ALERT: mikro04;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 75.60 ms
[1541645117] SERVICE ALERT: mikro04;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 63.71 ms
[1541646073] SERVICE FLAPPING ALERT: mikro52;Ping;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1541646131] SERVICE FLAPPING ALERT: mikro61;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541646189] SERVICE FLAPPING ALERT: mikro69;Ping;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1541646435] SERVICE FLAPPING ALERT: mikro57;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541646589] SERVICE FLAPPING ALERT: mikro48;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541646928] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.75 ms
[1541646937] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1541646945] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541646949] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541647009] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.79 ms
[1541647013] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.88 ms
[1541647188] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541647225] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541647326] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541647683] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541647769] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1541647806] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541647808] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1541647840] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.29 ms
[1541647844] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.22 ms
[1541647879] HOST ALERT: mikro41;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541647916] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541647944] HOST ALERT: mikro41;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 12.27 ms
[1541647958] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1541648016] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541648018] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541648022] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541648078] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.71 ms
[1541648082] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.34 ms
[1541648138] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541648139] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.41 ms
[1541648148] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541648170] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541648179] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.68 ms
[1541648203] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541648299] SERVICE ALERT: mikro41;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541648349] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.65 ms
[1541648423] SERVICE ALERT: mikro41;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 115.83 ms
[1541648494] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541648783] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541648787] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541648962] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.92 ms
[1541648986] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.81 ms
[1541649060] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541649064] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541649098] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.87 ms
[1541649102] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1541649105] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 3.16 ms
[1541649130] SERVICE ALERT: mikro61;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541649254] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541649398] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1541650053] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541650087] HOST ALERT: nev-k4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541650176] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541650180] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541650689] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541650771] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541650774] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541650965] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541651081] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.82 ms
[1541651085] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.37 ms
[1541651167] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541651175] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.57 ms
[1541652523] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541652571] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541652601] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541652615] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541652619] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541652802] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541652886] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1541652948] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.22 ms
[1541653204] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.06 ms
[1541653371] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1541653880] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.67 ms
[1541653884] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.80 ms
[1541654170] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.28 ms
[1541654384] HOST ALERT: mikro42;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.42.200)
[1541654448] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.01 ms
[1541654456] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.25 ms
[1541654495] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 261.85 ms
[1541654499] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541654825] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541654829] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541654852] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.90 ms
[1541654938] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1541654962] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.97 ms
[1541654984] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.19 ms
[1541655041] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.22 ms
[1541655080] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1541655084] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1541655085] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.75 ms
[1541655089] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541655322] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 10.32 ms
[1541655326] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.35 ms
[1541655525] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.51 ms
[1541655560] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.58 ms
[1541655564] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.70 ms
[1541655647] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.36 ms
[1541655746] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 16.72 ms
[1541655750] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.65 ms
[1541655860] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.52 ms
[1541655938] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.37 ms
[1541656100] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541656231] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1541656249] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1541656395] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.40 ms
[1541656408] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541656497] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541656500] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541656553] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541656739] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541656766] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541656771] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1541656804] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.70 ms
[1541656808] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1541656939] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.35 ms
[1541657003] HOST ALERT: sad103-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541657038] SERVICE ALERT: mikro43;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.200)
[1541657067] HOST ALERT: sad103-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1541657162] SERVICE ALERT: mikro43;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.90 ms
[1541657224] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 29.30 ms
[1541660422] SERVICE ALERT: mikro41;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541660546] SERVICE ALERT: mikro41;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.22 ms
[1541660571] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541660873] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541662023] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541662261] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.31 ms
[1541671690] SERVICE ALERT: kirochn-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 139.29 ms
[1541671791] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541671814] SERVICE ALERT: kirochn-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 78.52 ms
[1541671855] HOST ALERT: mikro70;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541672890] SERVICE ALERT: mikro27;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541673014] SERVICE ALERT: mikro27;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541673076] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541673200] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.90 ms
[1541674794] SERVICE ALERT: mikro44;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.200)
[1541674918] SERVICE ALERT: mikro44;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.91 ms
[1541675795] SERVICE ALERT: mikro28;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541675919] SERVICE ALERT: mikro28;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541676429] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541676494] HOST ALERT: nevsky130-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.18 ms
[1541677867] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541677870] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541678018] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541678050] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541679018] SERVICE ALERT: mikro36;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.36.200)
[1541679142] SERVICE ALERT: mikro36;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.13 ms
[1541687500] SERVICE ALERT: mikro52;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541687501] SERVICE ALERT: mikro54;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541687502] HOST ALERT: mikro53;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541687566] HOST ALERT: mikro53;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541687624] SERVICE ALERT: mikro52;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541687625] SERVICE ALERT: mikro54;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 14.20 ms
[1541692847] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541692973] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541693080] SERVICE ALERT: lunach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541693083] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541693147] HOST ALERT: lunach-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 11.85 ms
[1541693169] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541693187] HOST ALERT: mikro40;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1541693204] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 55.80 ms
[1541693251] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.61 ms
[1541693584] HOST ALERT: mikro60;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541693624] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541694400] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541694519] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541694527] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 204.31 ms
[1541694569] SERVICE ALERT: sad29-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541694572] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541694634] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541694692] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541694698] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541694814] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541694845] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541695453] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541695580] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1541695580] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541695906] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541695965] HOST ALERT: mikro57;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541695971] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.07 ms
[1541696029] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 11.09 ms
[1541696064] HOST ALERT: nev6-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541696190] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1541696270] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541696321] SERVICE ALERT: mikro63;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541696334] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541696359] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541696399] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541696445] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.25 ms
[1541696445] HOST ALERT: mikro42;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.42.200)
[1541696463] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.75 ms
[1541696509] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.28 ms
[1541696789] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541696961] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1541696970] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1541697043] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.53 ms
[1541697141] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541697270] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1541697270] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541697647] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541697973] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541698153] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541698247] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541698573] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541698753] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541698793] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541698796] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541698848] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 4.05 ms
[1541698917] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.86 ms
[1541698921] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 9.73 ms
[1541699220] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541699284] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541699773] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541699789] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (30.9% change > 20.0% threshold)
[1541699789] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (31.2% change > 20.0% threshold)
[1541699892] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (30.9% change > 20.0% threshold)
[1541699892] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (30.4% change > 20.0% threshold)
[1541699954] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1541700047] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541700144] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541700174] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700222] SERVICE ALERT: garkav-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541700252] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700267] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541700296] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541700300] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541700373] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541700373] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700374] HOST ALERT: engels-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700392] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541700398] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (25.6% change > 20.0% threshold)
[1541700398] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (39.6% change > 20.0% threshold)
[1541700401] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541700422] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541700423] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541700426] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541700430] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700432] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541700491] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541700500] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541700518] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541700522] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541700527] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541700539] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541700546] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541700552] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541700553] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541700556] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541700570] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541700570] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541700603] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541700614] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541700617] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541700625] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541700646] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541700646] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541700648] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541700648] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.11 ms
[1541700655] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541700671] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541700679] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541700693] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (24.9% change > 20.0% threshold)
[1541700693] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (35.1% change > 20.0% threshold)
[1541700704] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541700725] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541700737] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541700743] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541700759] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541700772] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541700772] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541700799] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541700830] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541700863] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541700907] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541700921] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541700937] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541700949] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541700963] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541700989] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541700994] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541701003] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541701033] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541701078] HOST ALERT: tur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541701089] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541701159] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541701161] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541701203] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541701204] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541701259] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541701296] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541701328] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541701454] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541701549] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541701573] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541701624] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541701675] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541701701] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541701738] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541701753] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541701827] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541701847] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541701879] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541701882] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541701970] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541702002] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541702005] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541702173] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541702230] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541702304] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541702353] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541702356] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541702447] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541702512] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541702638] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541702703] HOST FLAPPING ALERT: dek-k2;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541702773] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541702808] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541702898] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541702953] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541703003] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541703070] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541703078] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.02 seconds)
[1541703194] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1541703194] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541703198] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541703879] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541704001] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541704002] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541704004] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541704124] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541704130] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541704175] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541704268] HOST ALERT: prosv87-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541704394] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541704421] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541704424] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541704478] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541704481] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541704485] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541704488] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541704546] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541704601] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541704604] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541704604] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541704608] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541704611] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541704724] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541704848] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541704972] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541705088] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541705145] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541705208] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541705214] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541705385] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541705745] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541705834] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541706014] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541706285] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541706305] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541706349] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.81 ms
[1541706411] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541706449] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541706575] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541706675] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541706745] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541706829] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.81 ms
[1541706947] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541707098] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541707222] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1541707277] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541707457] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541707536] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541707554] HOST ALERT: msk197-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541707573] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541707576] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541707618] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541707662] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541707696] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541707699] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541707707] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541707744] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541707787] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541707824] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541707827] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541707867] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541707897] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541707943] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541707993] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541708043] SERVICE ALERT: nev-k7;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541708073] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708095] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708194] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541708249] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708253] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541708275] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541708328] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708429] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541708453] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708483] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541708488] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541708508] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541708518] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 18.03 ms
[1541708543] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541708843] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541708998] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.86 ms
[1541709323] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541709504] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541709649] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541709774] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.25 ms
[1541710078] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541710974] SERVICE ALERT: mikro56;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.52 ms
[1541711005] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541711098] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.27 ms
[1541711102] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1541711428] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541711493] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 6.49 ms
[1541711725] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541711849] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.20 ms
[1541712055] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541712120] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 4.26 ms
[1541712300] HOST ALERT: hoshimin-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541712301] HOST ALERT: hoshimin-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541712327] SERVICE ALERT: hoshimin-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541712364] HOST ALERT: hoshimin-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 87.14 ms
[1541712365] HOST ALERT: hoshimin-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 47.92 ms
[1541712445] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541712451] SERVICE ALERT: hoshimin-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 83.88 ms
[1541712453] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 20.92 ms
[1541713379] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1541713534] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.51 ms
[1541713832] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541713896] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 21.34 ms
[1541714222] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541714276] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541714378] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541714400] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541714404] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1541715030] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541715095] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 8.82 ms
[1541716620] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541716800] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541717403] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541717528] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.12 ms
[1541717832] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.14 ms
[1541718458] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541718638] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541718731] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541719030] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541719095] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.08 ms
[1541719720] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541719900] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541719954] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541720078] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541720082] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1541720408] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541720472] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1541720705] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541720735] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541720800] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.18 ms
[1541720829] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541721125] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541721252] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541721456] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541721486] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541721551] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.19 ms
[1541721580] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms