[1541121469] SERVICE ALERT: mikro44;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.200)
[1541121593] SERVICE ALERT: mikro44;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.80 ms
[1541128979] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.72 ms
[1541128983] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.86 ms
[1541129035] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.65 ms
[1541129039] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541129069] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541129073] HOST ALERT: nev6-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541129084] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541129085] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541129094] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541129113] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541129137] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541129148] HOST ALERT: nev-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541129159] SERVICE ALERT: nev-k9;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 5.43 ms
[1541129164] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1541129168] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.52 ms
[1541129209] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1541129252] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541129310] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541129340] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541129358] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1541129362] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541129404] HOST ALERT: nev-k9;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541129426] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541129434] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541129457] SERVICE ALERT: nev6-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541129484] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541129488] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541129508] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.75 ms
[1541129513] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541129604] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541129628] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541129646] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.97 ms
[1541129850] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541129854] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541129937] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541129962] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.75 ms
[1541130086] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541130150] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541130242] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541130269] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541130510] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1541130541] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.92 ms
[1541130545] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.64 ms
[1541130682] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541130686] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541131684] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541131686] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.52 ms
[1541131689] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541131857] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541131865] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1541131917] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541131948] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541131952] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541132109] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1541132120] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.94 ms
[1541132152] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.74 ms
[1541132226] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541132230] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1541132275] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541132516] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.17 ms
[1541132520] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.42 ms
[1541133480] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1541133484] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541133682] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1541133850] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1541134294] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.96 ms
[1541134384] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.85 ms
[1541134850] HOST ALERT: mikro42;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.42.200)
[1541134866] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.77 ms
[1541134914] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.05 ms
[1541134930] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.28 ms
[1541134962] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.69 ms
[1541135123] HOST ALERT: mikro25;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.200)
[1541135166] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1541135186] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541135187] HOST ALERT: mikro25;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1541135190] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541135256] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.81 ms
[1541135346] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.16 ms
[1541135350] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1541135356] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541136869] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.40 ms
[1541136937] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.48 ms
[1541136966] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1541136986] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.31 ms
[1541136992] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1541137115] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541137118] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1541137163] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1541137256] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.71 ms
[1541137264] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.66 ms
[1541137310] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1541137365] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541137399] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.10 ms
[1541137571] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541137582] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.22 ms
[1541137598] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1541137605] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.97 ms
[1541137623] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.50 ms
[1541137628] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541137638] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.04 ms
[1541137823] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 14.55 ms
[1541137826] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.97 ms
[1541137923] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1541138079] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1541138084] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541138088] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.89 ms
[1541138094] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.91 ms
[1541138130] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1541138190] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.61 ms
[1541138195] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.19 ms
[1541138199] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.01 ms
[1541138374] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.69 ms
[1541139026] HOST ALERT: mikro28;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.200)
[1541139052] HOST ALERT: mikro43;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.200)
[1541139090] HOST ALERT: mikro28;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541139116] HOST ALERT: mikro43;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541139829] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1541140120] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541141733] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541141736] SERVICE ALERT: mikro27;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541141860] SERVICE ALERT: mikro27;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1541142014] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541142800] SERVICE ALERT: mikro72;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541142924] SERVICE ALERT: mikro72;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.02 ms
[1541145323] SERVICE ALERT: mikro68;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.68.200)
[1541145447] SERVICE ALERT: mikro68;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.21 ms
[1541145573] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.90 ms
[1541145858] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541153696] SERVICE ALERT: mikro69;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541153820] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.78 ms
[1541155548] HOST ALERT: mikro41;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541155565] SERVICE ALERT: nevsky78-k3;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 6.89 ms
[1541155612] HOST ALERT: mikro41;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 16.63 ms
[1541155689] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541161577] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541161641] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.71 ms
[1541162532] SERVICE ALERT: mikro25;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.200)
[1541162536] HOST ALERT: mikro25;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.200)
[1541162600] HOST ALERT: mikro25;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.09 ms
[1541162656] SERVICE ALERT: mikro25;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.78 ms
[1541164376] HOST ALERT: mikro61;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541164440] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.17 ms
[1541165310] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541165374] HOST ALERT: nevsky78-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.76 ms
[1541166250] SERVICE ALERT: mikro41;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541166374] SERVICE ALERT: mikro41;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.58 ms
[1541166659] HOST ALERT: mikro46;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.200)
[1541166709] HOST ALERT: mikro46;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1541170540] HOST ALERT: mikro36;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.36.200)
[1541170604] HOST ALERT: mikro36;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.30 ms
[1541171745] HOST ALERT: mikro39;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.200)
[1541171809] HOST ALERT: mikro39;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.95 ms
[1541174472] HOST ALERT: mikro54;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.200)
[1541174536] HOST ALERT: mikro54;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.49 ms
[1541174552] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541174631] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541174676] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541174746] HOST ALERT: mikro40;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1541174810] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.37 ms
[1541176165] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541176345] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541176462] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541177068] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541177069] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541177229] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541177229] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541177248] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541177355] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541178951] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541179074] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541179074] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541179141] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541179267] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1541179452] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541181278] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541181329] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541181404] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541181723] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541181743] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541181849] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541181886] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541181889] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541181911] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541181914] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541181970] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541181991] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541181993] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541181997] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541182004] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541182009] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541182015] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541182016] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182027] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541182034] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541182037] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541182093] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541182114] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541182114] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541182119] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541182124] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541182169] HOST ALERT: garkav-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.21)
[1541182172] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541182263] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182281] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541182281] HOST ALERT: mikro59;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.200)
[1541182300] HOST ALERT: mikro59;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541182308] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541182312] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182322] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541182329] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541182361] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541182383] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541182389] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541182407] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541182438] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541182444] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541182448] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541182449] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541182455] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541182468] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541182493] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541182570] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541182590] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182594] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541182620] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541182622] SERVICE ALERT: tur-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541182623] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541182625] HOST ALERT: tur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182665] SERVICE ALERT: engels-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541182695] HOST ALERT: engels-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541182716] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541182717] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541182738] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541182743] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541182745] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541182748] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541182749] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541182766] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541182788] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541182821] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541182983] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541182998] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541183109] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541183163] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541183279] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541183289] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541183314] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541183355] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541183488] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541183494] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541184042] SERVICE ALERT: nev-k11;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541184060] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541184193] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541184240] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541184244] SERVICE ALERT: msk197-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541184263] HOST ALERT: msk197-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541184367] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541184387] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541184665] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541184791] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541184985] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541185055] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541185111] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541185274] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541185282] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541185444] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541185447] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541185454] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541185508] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541185519] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541185567] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541185570] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541185585] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541185588] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541185634] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541185637] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541185708] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541185711] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541185745] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541185748] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541185868] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541185874] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541185947] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541185950] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541186070] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541186073] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541187874] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541188000] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541188195] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541188311] SERVICE ALERT: kamen-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541188314] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541188375] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541188378] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541188434] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541188437] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541188737] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541188801] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541188927] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541188953] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541189058] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541189066] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541189079] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541189167] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541189192] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541189213] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541189249] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541189286] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541189294] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541189339] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541189357] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541189366] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541189420] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541189453] SERVICE ALERT: nev-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541189459] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541189463] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541189574] HOST ALERT: mikro44;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.200)
[1541189601] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541189604] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541189631] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541189638] HOST ALERT: mikro44;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1541189639] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541189643] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541189760] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541190031] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541190061] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541190211] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541193199] HOST ALERT: mikro25;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.200)
[1541193260] HOST ALERT: mikro25;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.90 ms
[1541196255] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541196381] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541196519] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)