[1540980992] SERVICE ALERT: sadov-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.14.21)
[1540980995] HOST ALERT: sadov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.14.21)
[1540980998] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1540981012] SERVICE ALERT: sadov-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.14.25)
[1540981014] SERVICE ALERT: nev-k6;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.26)
[1540981015] HOST ALERT: sadov-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.14.25)
[1540981017] HOST ALERT: nev-k6;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.26)
[1540981040] SERVICE ALERT: nev-k8;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.28)
[1540981043] HOST ALERT: nev-k8;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.28)
[1540981067] HOST ALERT: karp10-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981068] HOST ALERT: mikro11;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.200)
[1540981078] SERVICE ALERT: karp10-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981080] SERVICE ALERT: pion12-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.53.22)
[1540981081] SERVICE ALERT: nev-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.24)
[1540981083] HOST ALERT: pion12-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.53.22)
[1540981084] HOST ALERT: nev-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.24)
[1540981086] HOST ALERT: michur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981088] SERVICE ALERT: sad103-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.60.23)
[1540981089] HOST ALERT: michur-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981091] HOST ALERT: sad103-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.60.23)
[1540981093] HOST ALERT: mikro08;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981097] HOST ALERT: mikro15;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981106] HOST ALERT: mikro26;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981107] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.29)
[1540981110] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.22)
[1540981110] HOST ALERT: nev-k9;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.29)
[1540981112] HOST ALERT: mikro29;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981113] HOST ALERT: mikro45;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.45.200)
[1540981113] HOST ALERT: nev-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.22)
[1540981115] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.14.21)
[1540981118] HOST ALERT: mikro34;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981121] HOST ALERT: sadov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.14.21)
[1540981123] SERVICE ALERT: michur-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981124] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1540981131] SERVICE ALERT: sadov-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.14.22)
[1540981132] SERVICE ALERT: nev-k11;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.33)
[1540981134] HOST ALERT: sadov-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.14.22)
[1540981135] HOST ALERT: nev-k11;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.33)
[1540981135] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.14.25)
[1540981137] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.26)
[1540981139] HOST ALERT: sadov-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.14.25)
[1540981143] HOST ALERT: nev-k6;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.26)
[1540981149] HOST ALERT: mikro67;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.200)
[1540981153] SERVICE ALERT: mikro71;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.71.200)
[1540981156] HOST ALERT: mikro71;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.200)
[1540981157] HOST ALERT: mikro55;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981163] SERVICE ALERT: nev-k8;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.28)
[1540981166] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.15.21)
[1540981166] HOST ALERT: nev-k8;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.28)
[1540981169] HOST ALERT: nev-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.21)
[1540981172] HOST ALERT: mikro62;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981174] HOST ALERT: nev-k10;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.32)
[1540981181] HOST ALERT: nev-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.23)
[1540981185] HOST ALERT: nev-k5;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.25)
[1540981189] HOST ALERT: nev-k7;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.15.27)
[1540981194] HOST ALERT: mikro11;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.200)
[1540981195] HOST ALERT: nastav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981197] HOST ALERT: nastav-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981203] SERVICE ALERT: pion12-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.53.22)
[1540981204] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.24)
[1540981209] HOST ALERT: pion12-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.53.22)
[1540981210] HOST ALERT: nev-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.24)
[1540981211] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.23)
[1540981211] SERVICE ALERT: sad103-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.60.23)
[1540981217] HOST ALERT: sad103-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.60.23)
[1540981226] SERVICE ALERT: sadov-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.14.30)
[1540981229] HOST ALERT: sadov-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.14.30)
[1540981230] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.29)
[1540981233] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.22)
[1540981236] SERVICE ALERT: sadov-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (172.17.14.27)
[1540981236] HOST ALERT: nev-k9;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.29)
[1540981239] HOST ALERT: sadov-k5;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.17.14.27)
[1540981239] HOST ALERT: nev-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.22)
[1540981239] HOST ALERT: mikro45;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.200)
[1540981247] HOST ALERT: karp10-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981254] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.14.22)
[1540981255] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.33)
[1540981260] HOST ALERT: sadov-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.14.22)
[1540981261] HOST ALERT: nev-k11;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.33)
[1540981266] HOST ALERT: michur-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981269] HOST ALERT: michur-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981273] HOST ALERT: mikro08;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981275] HOST ALERT: mikro67;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.200)
[1540981276] SERVICE ALERT: mikro71;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.200)
[1540981277] HOST ALERT: mikro15;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981282] SERVICE ALERT: mikro67;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.200)
[1540981282] HOST ALERT: mikro71;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.200)
[1540981286] HOST ALERT: mikro26;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981289] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.21)
[1540981292] HOST ALERT: mikro29;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981295] HOST ALERT: nev-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.21)
[1540981297] HOST ALERT: yaroslav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981298] HOST ALERT: mikro34;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981299] HOST ALERT: yaroslav-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540981300] HOST ALERT: nev-k10;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.32)
[1540981307] HOST ALERT: nev-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.23)
[1540981311] HOST ALERT: nev-k5;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.25)
[1540981311] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.25)
[1540981315] HOST ALERT: nev-k7;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.15.27)
[1540981337] HOST ALERT: mikro55;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981349] SERVICE ALERT: sadov-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.14.30)
[1540981352] HOST ALERT: mikro62;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981355] HOST ALERT: sadov-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.14.30)
[1540981356] SERVICE ALERT: mikro11;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.200)
[1540981359] SERVICE ALERT: sadov-k5;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.14.27)
[1540981360] SERVICE ALERT: yaroslav-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981365] HOST ALERT: sadov-k5;DOWN;HARD;3;CRITICAL - Host Unreachable (172.17.14.27)
[1540981375] HOST ALERT: nastav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981377] HOST ALERT: nastav-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981380] SERVICE ALERT: mikro45;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.200)
[1540981476] SERVICE ALERT: mikro34;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981477] HOST ALERT: yaroslav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981479] HOST ALERT: yaroslav-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1540981481] SERVICE ALERT: nev-k10;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.32)
[1540981511] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (172.17.15.27)
[1540981538] SERVICE ALERT: yaroslav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981544] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1540981576] SERVICE ALERT: mikro29;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540981600] SERVICE ALERT: michur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540984321] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1540984385] HOST ALERT: gavan-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 9.12 ms
[1540984665] HOST ALERT: mikro54;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.200)
[1540984729] HOST ALERT: mikro54;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 35.78 ms
[1540987628] HOST ALERT: bkuna-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987629] HOST ALERT: domino;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987630] HOST ALERT: blago-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987630] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987631] HOST ALERT: dragon;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987632] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987633] HOST ALERT: eliz-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987635] HOST ALERT: eliz-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987638] HOST ALERT: engels-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987640] SERVICE ALERT: metal-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987640] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987642] HOST ALERT: esx2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987643] HOST ALERT: mikro04;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987643] HOST ALERT: metal-k1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1540987645] SERVICE ALERT: mikro31;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987646] SERVICE ALERT: mikro37;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987646] HOST ALERT: gasheka-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987647] SERVICE ALERT: lunach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987648] HOST ALERT: mikro31;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987649] HOST ALERT: mikro37;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987650] HOST ALERT: glassfish;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987650] HOST ALERT: lunach-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987652] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987652] HOST ALERT: mikro203;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987652] HOST ALERT: grivcov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987654] HOST ALERT: mikro25;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987654] HOST ALERT: lensovet-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.32.21)
[1540987654] SERVICE ALERT: mikro39;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.200)
[1540987654] HOST ALERT: mikro47;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.200)
[1540987654] SERVICE ALERT: varsh-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1540987654] SERVICE ALERT: dach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.52.21)
[1540987654] HOST ALERT: mikro60;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.60.200)
[1540987654] HOST ALERT: mikro48;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.200)
[1540987654] HOST ALERT: mikro46;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.200)
[1540987654] HOST ALERT: mikro39;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.200)
[1540987654] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1540987654] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1540987654] HOST ALERT: lion;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.16.0.33)
[1540987655] HOST ALERT: hoshimin-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987655] SERVICE ALERT: nev6-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987655] HOST ALERT: kamen-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987656] SERVICE ALERT: mikro27;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987657] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1540987657] HOST ALERT: dach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.52.21)
[1540987657] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1540987657] HOST ALERT: mikro49;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.200)
[1540987657] HOST ALERT: hoshimin-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987658] HOST ALERT: nev6-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987658] HOST ALERT: mikro28;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1540987659] HOST ALERT: intra;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987659] HOST ALERT: mikro27;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987660] SERVICE ALERT: mikro61;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1540987660] HOST ALERT: mikro52;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.52.200)
[1540987660] SERVICE ALERT: mikro69;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1540987660] HOST ALERT: mikro32;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987661] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987662] SERVICE ALERT: dragon;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540987662] HOST ALERT: mikro35;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987663] HOST ALERT: mikro53;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.53.200)
[1540987663] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1540987663] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1540987664] HOST ALERT: mikro36;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987664] SERVICE ALERT: mikro24;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1540987665] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987666] SERVICE ALERT: mikro40;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1540987666] SERVICE ALERT: mikro48;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.200)
[1540987666] HOST ALERT: mikro40;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987667] HOST ALERT: mikro24;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987669] HOST ALERT: mikro41;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987669] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1540987669] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1540987669] HOST ALERT: kirochn-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987671] HOST ALERT: mikro42;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987672] HOST ALERT: mikro58;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.200)
[1540987672] HOST ALERT: kirochn-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987673] HOST ALERT: mikro43;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987674] SERVICE ALERT: mikro60;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540987674] HOST ALERT: kronsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987675] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1540987675] HOST ALERT: mikro59;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.200)
[1540987675] HOST ALERT: mikro44;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987676] HOST ALERT: lens22-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540987682] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1540987682] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1540987685] HOST ALERT: mikro68;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.200)
[1540987689] SERVICE ALERT: mikro49;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.200)
[1540987689] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1540987689] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1540987692] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1540987695] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1540987695] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1540987698] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1540987698] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1540987698] HOST ALERT: nas;DOWN;SOFT;1;CRITICAL - Host Unreachable (172.16.0.4)
[1540987698] SERVICE ALERT: kirochn-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.22)
[1540987701] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1540987701] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1540987704] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1540987704] SERVICE ALERT: mikro46;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.200)
[1540987710] SERVICE ALERT: mikro68;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.200)
[1540987720] HOST ALERT: nev6-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1540987720] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1540987723] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1540987723] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1540987726] HOST ALERT: mikro54;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.200)
[1540987726] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1540987729] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1540991751] HOST ALERT: mikro68;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1540991751] HOST ALERT: mikro46;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.14 ms
[1540991752] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1540991752] HOST ALERT: mikro59;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1540991753] HOST ALERT: 8line-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1540991754] HOST ALERT: lion;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1540991754] HOST ALERT: varsh-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1540991755] SERVICE ALERT: dach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.77 ms
[1540991756] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1540991757] HOST ALERT: aviakonstr-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.71 ms
[1540991757] HOST ALERT: nevsky130-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1540991757] HOST ALERT: mikro43;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1540991758] HOST ALERT: nev6-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1540991759] HOST ALERT: mikro47;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1540991759] HOST ALERT: dach-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.40 ms
[1540991760] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1540991760] HOST ALERT: mikro37;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1540991761] HOST ALERT: mikro32;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.29 ms
[1540991761] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.94 ms
[1540991761] HOST ALERT: bkuna-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 10.08 ms
[1540991762] HOST ALERT: mikro53;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1540991762] HOST ALERT: mikro58;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1540991762] HOST ALERT: mikro69;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.11 ms
[1540991762] HOST ALERT: mikro31;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.27 ms
[1540991763] HOST ALERT: blago-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1540991763] HOST ALERT: lunach-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 92.57 ms
[1540991765] HOST ALERT: bryanc20-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1540991765] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.93 ms
[1540991765] HOST ALERT: eliz-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1540991765] HOST ALERT: metal-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.02 ms
[1540991766] HOST ALERT: mikro24;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.29 ms
[1540991766] HOST ALERT: mikro39;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1540991766] HOST ALERT: kirochn-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 52.77 ms
[1540991767] HOST ALERT: kronsh-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.94 ms
[1540991767] HOST ALERT: nevsk104-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1540991768] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1540991768] HOST ALERT: mikro35;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1540991770] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 68.30 ms
[1540991771] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1540991771] HOST ALERT: nevsk104-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1540991772] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1540991772] HOST ALERT: mikro36;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.53 ms
[1540991773] HOST ALERT: msk197-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1540991776] HOST ALERT: domino;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1540991776] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.38 ms
[1540991777] HOST ALERT: nevsky130-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1540991778] HOST ALERT: dragon;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1540991778] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1540991779] HOST ALERT: mikro48;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1540991779] HOST ALERT: mikro52;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1540991779] HOST ALERT: mikro49;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1540991779] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.45 ms
[1540991780] HOST ALERT: eliz-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1540991780] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 7.38 ms
[1540991780] HOST ALERT: mikro04;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 64.59 ms
[1540991781] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1540991782] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1540991785] HOST ALERT: engels-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.98 ms
[1540991786] HOST ALERT: nevsk104-k4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.89 ms
[1540991787] HOST ALERT: mikro25;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.61 ms
[1540991787] HOST ALERT: engels51-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.38 ms
[1540991788] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.76 ms
[1540991788] HOST ALERT: kamen-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1540991789] HOST ALERT: esx2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1540991789] HOST ALERT: mikro28;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1540991791] HOST ALERT: garkav-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.40 ms
[1540991791] HOST ALERT: nev6-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1540991791] HOST ALERT: kirochn-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 72.61 ms
[1540991793] HOST ALERT: gasheka-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.98 ms
[1540991793] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.90 ms
[1540991794] HOST FLAPPING ALERT: mikro54;STARTED; Host appears to have started flapping (20.5% change > 20.0% threshold)
[1540991794] HOST ALERT: mikro54;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 14.72 ms
[1540991795] HOST ALERT: mikro70;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1540991795] HOST FLAPPING ALERT: gavan-k1;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1540991795] HOST ALERT: gavan-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1540991796] SERVICE ALERT: mikro39;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.38 ms
[1540991796] HOST ALERT: mikro27;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1540991797] HOST ALERT: glassfish;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1540991798] HOST ALERT: mikro203;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.43 ms
[1540991798] HOST ALERT: mikro41;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 15.71 ms
[1540991799] HOST ALERT: nas;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1540991799] HOST ALERT: grivcov-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1540991802] HOST ALERT: hoshimin-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 86.99 ms
[1540991803] HOST ALERT: lensovet-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.05 ms
[1540991804] HOST ALERT: hoshimin-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 25.41 ms
[1540991806] HOST ALERT: intra;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1540991806] HOST ALERT: nevsk104-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1540991808] HOST ALERT: mikro44;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1540991808] HOST ALERT: msk197-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.82 ms
[1540991808] HOST ALERT: kamen-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1540991810] HOST ALERT: kamen-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1540991811] SERVICE ALERT: mikro24;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.70 ms
[1540991814] SERVICE ALERT: mikro37;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1540991814] SERVICE ALERT: mikro31;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.61 ms
[1540991816] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1540991843] SERVICE ALERT: mikro40;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.99 ms
[1540991844] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.73 ms
[1540991852] SERVICE ALERT: mikro27;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1540991861] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1540991866] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.89 ms
[1540991866] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1540991877] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1540991908] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1540991912] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.91 ms
[1540991928] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1540991932] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1540991935] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1540991939] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1540991946] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1540991951] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1540991953] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1540991957] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1540991959] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1540991963] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1540991964] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1540991975] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1540991985] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.92 ms
[1540991989] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1540991996] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1540992000] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1540992016] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1540992018] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1540992025] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1540992073] SERVICE ALERT: dragon;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1540992096] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1540992096] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.91 ms
[1540992101] SERVICE ALERT: mikro46;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1540992214] HOST FLAPPING ALERT: mikro54;STARTED; Host appears to have started flapping (20.0% change > 20.0% threshold)
[1540992238] SERVICE ALERT: mikro49;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1540992251] SERVICE ALERT: mikro60;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1540992257] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.74 ms
[1540992266] SERVICE ALERT: mikro68;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1540992277] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1540992336] HOST ALERT: turist28-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.21)
[1540992338] SERVICE ALERT: kirochn-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 155.56 ms
[1540992343] SERVICE ALERT: mikro48;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1540992400] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.23 ms
[1540992575] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1540993561] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1540993570] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1540993617] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1540993625] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1540993625] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1540993634] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1540993681] HOST FLAPPING ALERT: garkav-k1;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1540993681] HOST ALERT: garkav-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.79 ms
[1540994051] SERVICE ALERT: mikro203;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 9.70 ms
[1540994175] SERVICE ALERT: mikro203;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1540995694] HOST FLAPPING ALERT: mikro54;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1540996279] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1540996343] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.69 ms
[1540997500] HOST FLAPPING ALERT: garkav-k1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1540997525] HOST FLAPPING ALERT: mikro56;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541001544] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541001547] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541001667] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541001673] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541005228] SERVICE ALERT: sadov-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541005352] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1541005542] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541005722] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541005927] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541008901] SERVICE ALERT: garkav-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 3.85 ms
[1541008931] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541009052] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541009109] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541009111] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541009157] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541009158] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541009162] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541009164] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541009169] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541009170] SERVICE ALERT: nev6-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541009174] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541009184] HOST ALERT: nev6-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541009192] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541009221] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541009235] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541009254] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541009260] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541009263] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541009267] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541009270] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541009283] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541009288] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541009290] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541009293] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541009300] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541009310] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1541009317] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541009338] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541009348] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541009383] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541009386] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541009387] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541009390] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541009394] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541009457] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541009468] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541009576] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541009579] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541009583] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541009594] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541009760] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541009760] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541009785] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541009794] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541009818] SERVICE ALERT: mikro57;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 6.34 ms
[1541009865] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541009887] SERVICE ALERT: lens22-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541009890] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541009911] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541009920] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541009927] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541009942] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.02 ms
[1541010010] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541010016] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541010030] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541010080] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541010092] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541010098] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541010149] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541010156] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541010206] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541010215] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541010224] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541010367] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541010443] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541010485] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541010626] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541010664] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541010752] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541010790] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541010926] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541010960] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541011224] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541011229] SERVICE ALERT: eliz-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541011232] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541011350] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541011352] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541011355] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541011387] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.15 ms
[1541011540] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541011583] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541011686] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541011709] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541011764] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541011812] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541012010] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541012309] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541012361] SERVICE ALERT: mikro63;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541012435] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541012485] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.32 ms
[1541012711] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541012727] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541012735] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541012837] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541012861] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541013274] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541013579] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541013601] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541013649] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541013652] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541013688] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541013700] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541013707] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541013727] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541013759] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541013778] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541013833] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541013860] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541014177] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541014198] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541014256] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541014324] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541014357] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541014520] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541014634] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541014724] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541014761] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541015115] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541015241] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541015350] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541015965] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541015966] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541015985] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541016065] HOST ALERT: mikro40;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.200)
[1541016127] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541016129] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.73 ms
[1541016145] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541016166] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541016228] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541016231] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541016272] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541016289] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541016351] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541016357] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541016372] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541016383] SERVICE ALERT: kamen-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541016386] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541016390] HOST ALERT: nev6-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1541016452] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541016469] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541016478] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541016498] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541016506] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541016508] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541016509] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541016509] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541016515] HOST ALERT: nev6-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.22)
[1541016595] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541016606] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541016631] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541016634] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541016636] SERVICE ALERT: nev6-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.22)
[1541017005] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541017131] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541017166] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541017266] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541017268] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541017391] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541017394] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541017483] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541017637] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541017663] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541017694] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541017806] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541029022] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541029086] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.54 ms