[1541980969] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541981086] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1541981086] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541981146] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1541981149] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541981243] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1541981243] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.52 ms
[1541981472] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541981569] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541981652] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541981747] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.84 ms
[1541981749] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541981843] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541982073] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541982169] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541982253] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541982346] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541982349] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541982443] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541982672] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541982769] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541982852] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541982896] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.15 ms
[1541982946] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541983222] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541983272] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541983287] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.25 ms
[1541983452] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541983546] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541983612] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1541983643] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541983872] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541983969] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541984000] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541984149] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541984243] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 23.95 ms
[1541984326] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541984506] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541984866] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541984869] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541984991] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.07 ms
[1541984995] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1541985047] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541985321] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541985369] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541985373] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541985493] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1541985497] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.07 ms
[1541985501] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541985594] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541985823] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1541986003] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541986097] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 26.68 ms
[1541986220] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541986375] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.60 ms
[1541986694] SERVICE ALERT: mikro56;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.13 ms
[1541986723] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541986817] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541986818] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 12.98 ms
[1541986822] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1541986941] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541987445] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541987447] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541987448] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541987545] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.07 ms
[1541987596] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541987628] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541988168] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541988169] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541988171] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541988173] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.15 ms
[1541988293] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.13 ms
[1541988297] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541988623] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541988799] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541988803] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541988896] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1541988954] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.27 ms
[1541989222] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541989252] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541989316] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 26.10 ms
[1541989396] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541989402] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541989496] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.03 ms
[1541989547] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541989671] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541989726] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541989907] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1541990122] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541990211] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541990277] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.79 ms
[1541990537] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541990717] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541990719] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541990844] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.21 ms
[1541990898] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1541990898] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541991023] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541991087] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541991321] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541991413] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541991447] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541991622] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541991705] SERVICE ALERT: mikro48;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541991708] HOST ALERT: mikro48;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541991773] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541991828] SERVICE ALERT: mikro48;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541991831] HOST ALERT: mikro48;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.200)
[1541991900] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541991947] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541992104] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541992153] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541992171] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1541992226] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541992226] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541992406] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541992462] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541992553] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541992598] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541992648] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 2.17 ms
[1541992733] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541992827] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541993128] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541993138] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541993142] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541993192] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541993253] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 52.29 ms
[1541993273] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541993323] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541993338] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.27 ms
[1541993378] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541993382] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1541993391] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541993396] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.69 ms
[1541993421] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.94 ms
[1541993425] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1541993433] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1541993443] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1541993447] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.63 ms
[1541993492] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541993547] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541993554] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541993556] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.88 ms
[1541993601] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541993679] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.94 ms
[1541993683] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541993824] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541993828] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541993838] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541993845] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541993848] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541993870] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541993883] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1541993888] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1541993920] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541993927] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541993928] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541993984] HOST ALERT: nev-k6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1541993995] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.12 ms
[1541994013] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.26 ms
[1541994034] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541994049] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541994052] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541994082] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1541994107] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541994152] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.20 ms
[1541994153] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541994169] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541994173] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541994200] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541994200] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541994324] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541994362] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.21 ms
[1541994364] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1541994379] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541994380] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541994559] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541994599] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541994647] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541994693] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541994696] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541994893] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.60 ms
[1541994897] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541994921] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541996687] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541996720] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.41 ms
[1541996925] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1541996957] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541996962] HOST FLAPPING ALERT: mikro56;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541997801] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541997899] HOST FLAPPING ALERT: dek-k1;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1541998006] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1541998010] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541998090] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1541998151] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1541998202] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.19 ms
[1541998206] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.19 ms
[1541998290] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541998319] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541998323] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541998379] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.50 ms
[1541998670] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541998752] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.09 ms
[1541998756] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 12.44 ms
[1541999077] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.76 ms
[1541999092] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.91 ms
[1541999606] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.68 ms
[1541999653] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.00 ms
[1541999722] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.94 ms
[1541999732] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541999901] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.75 ms
[1541999936] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.85 ms
[1541999985] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.63 ms
[1542000025] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1542000187] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1542000191] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1542000230] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1542000540] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.65 ms
[1542000973] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1542000977] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542001038] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1542001042] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1542001054] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1542001079] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542001147] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1542001419] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1542001432] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.39 ms
[1542001434] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.58 ms
[1542001441] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.43 ms
[1542001537] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 13.07 ms
[1542001760] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.66 ms
[1542001830] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.28 ms
[1542001854] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1542001914] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542001958] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.47 ms
[1542001962] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.79 ms
[1542001967] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542002050] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.52 ms
[1542002082] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1542002180] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1542002343] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1542002613] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1542002617] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542002673] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542002813] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542002833] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542002877] HOST ALERT: 8line-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1542002942] SERVICE ALERT: eliz-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542002945] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542002957] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1542003065] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542003071] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1542003375] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542003444] SERVICE FLAPPING ALERT: mikro56;Ping;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1542003666] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1542005029] SERVICE ALERT: mikro48;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 11.18 ms
[1542005033] HOST ALERT: mikro48;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.83 ms
[1542005102] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1542005106] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542005525] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1542005527] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1542009601] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542009665] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542009863] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.75 ms
[1542009991] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542010040] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542010045] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542010165] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542010171] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542010220] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542010221] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542010229] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1542010373] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.54 ms
[1542010555] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (24.1% change > 20.0% threshold)
[1542010555] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542010650] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 3.30 ms
[1542010775] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1542010775] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542011074] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542011138] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542011464] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542011546] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.40 ms
[1542011644] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542011875] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542011940] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.47 ms
[1542012150] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542012238] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542012302] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542012497] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542012628] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542012677] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542012771] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1542012771] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542012772] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542012808] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542012896] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542013049] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.12 ms
[1542013097] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542013102] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542013277] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542013282] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542013371] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.64 ms
[1542013375] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542013500] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542013555] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542013649] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542013697] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542013761] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1542013975] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542014123] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542014131] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 15.19 ms
[1542014247] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1542014424] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542014430] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542014488] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.06 ms
[1542014494] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542014566] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542014691] SERVICE ALERT: dek-k2;Ping;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 2.31 ms
[1542014721] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542014814] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542014815] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542014819] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.74 ms
[1542014820] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542014844] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542014846] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542014969] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.19 ms
[1542014969] SERVICE ALERT: mikro56;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 80%, RTA = 10.48 ms
[1542015000] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542015118] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542015182] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1542015569] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1542015569] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.97 ms
[1542015570] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1542015570] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542015573] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.72 ms
[1542015808] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542015876] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542015894] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542015940] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1542015970] SERVICE ALERT: sadov-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542015973] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542015988] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542016020] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.01 ms
[1542016037] HOST ALERT: sadov-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1542016074] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542016094] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1542016174] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1542016266] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542016346] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542016411] HOST ALERT: dek-k2;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 3.84 ms
[1542016446] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542016500] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542016642] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542016680] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542016766] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542016774] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542017050] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.23 ms
[1542017100] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542017280] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542017374] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542017376] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542017556] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542017573] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1542017700] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542017753] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542017880] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542017967] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542017971] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1542017974] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.04 ms
[1542018160] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542018277] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542018432] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.70 ms
[1542018691] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1542018761] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542018825] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542019293] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542019357] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542019444] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542019497] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542019512] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.51 ms
[1542019595] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1542019652] HOST ALERT: dek-k2;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.45 ms
[1542019923] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542020103] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542020199] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542020203] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542020506] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542020570] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.90 ms
[1542020753] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1542020753] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (36.8% change > 20.0% threshold)
[1542020753] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1542020845] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1542020845] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (35.9% change > 20.0% threshold)
[1542020845] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (42.1% change >= 20.0% threshold)
[1542020869] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542020919] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1542020919] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (41.2% change > 20.0% threshold)
[1542020919] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (42.1% change >= 20.0% threshold)
[1542020933] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1542020969] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1542020969] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (42.8% change > 20.0% threshold)
[1542020969] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (42.1% change >= 20.0% threshold)
[1542021001] HOST ALERT: michur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542021003] HOST ALERT: michur-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542021043] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1542021043] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (42.8% change > 20.0% threshold)
[1542021043] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (42.1% change >= 20.0% threshold)
[1542021177] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1542021177] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542021181] HOST ALERT: michur-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542021183] HOST ALERT: michur-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542021357] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542021600] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542021724] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542021930] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542022004] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542022110] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542022129] SERVICE ALERT: mikro56;Ping;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 2.33 ms
[1542022159] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542022182] HOST ALERT: bkuna-k1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542022246] HOST ALERT: bkuna-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.33 ms
[1542022251] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542022253] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542022257] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.20 ms
[1542022351] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542022415] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 60%, RTA = 2.70 ms
[1542022740] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542022846] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542022849] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542022851] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542022854] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542022913] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1542022918] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542022920] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542022970] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.31 ms
[1542022975] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 12.09 ms
[1542023239] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542023244] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542023419] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542023424] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542023525] HOST ALERT: dek-k1;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 4.94 ms
[1542023597] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542023602] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542023728] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542024024] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 8.39 ms
[1542024054] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542024150] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542024234] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542024275] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1542024275] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542024279] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542024321] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542024326] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542024330] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542024478] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (30.9% change > 20.0% threshold)
[1542024478] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (25.3% change > 20.0% threshold)
[1542024478] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (39.6% change > 20.0% threshold)
[1542024478] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1542024478] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1542024607] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (30.9% change > 20.0% threshold)
[1542024607] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (25.3% change > 20.0% threshold)
[1542024607] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (39.6% change > 20.0% threshold)
[1542024607] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1542024607] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1542024632] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542024649] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542024696] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 24.00 ms
[1542024803] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1542024925] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.50 ms
[1542025237] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542025322] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542025417] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542025440] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542025502] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542025502] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542025548] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542025551] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542025553] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542025564] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1542025699] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542025725] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 678.54 ms
[1542025731] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542026035] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.47 ms
[1542026047] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542026106] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542026112] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 3.32 ms
[1542026227] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.34 ms
[1542026272] SERVICE FLAPPING ALERT: dek-k2;Ping;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1542026272] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1542026277] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542026432] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542026437] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542026462] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (41.3% change > 20.0% threshold)
[1542026462] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (25.2% change > 20.0% threshold)
[1542026462] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (39.3% change > 20.0% threshold)
[1542026462] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1542026462] SERVICE FLAPPING ALERT: dek-k2;Ping;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1542026462] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (45.5% change >= 20.0% threshold)
[1542026587] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 9.73 ms
[1542026617] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542026661] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542026840] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (35.7% change > 20.0% threshold)
[1542026840] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (30.0% change > 20.0% threshold)
[1542026840] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (39.9% change > 20.0% threshold)
[1542026840] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1542026840] SERVICE FLAPPING ALERT: dek-k2;Ping;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1542026840] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (45.5% change >= 20.0% threshold)
[1542026844] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542026871] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542026874] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542026878] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542026938] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542026996] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.37 ms
[1542027029] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542027064] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (34.8% change > 20.0% threshold)
[1542027064] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (37.0% change > 20.0% threshold)
[1542027064] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (33.4% change > 20.0% threshold)
[1542027064] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1542027064] SERVICE FLAPPING ALERT: dek-k2;Ping;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1542027153] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (40.7% change >= 20.0% threshold)
[1542027153] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542027208] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542027221] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.45 ms
[1542027264] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542027388] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542027444] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1542027622] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542027846] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (34.7% change > 20.0% threshold)
[1542027846] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (35.6% change > 20.0% threshold)
[1542027846] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (24.5% change > 20.0% threshold)
[1542027846] SERVICE FLAPPING ALERT: dek-k1;Ping;STARTED; Service appears to have started flapping (26.4% change >= 20.0% threshold)
[1542027846] SERVICE FLAPPING ALERT: dek-k2;Ping;STARTED; Service appears to have started flapping (29.0% change >= 20.0% threshold)
[1542027846] SERVICE FLAPPING ALERT: mikro56;Ping;STARTED; Service appears to have started flapping (40.7% change >= 20.0% threshold)
[1542027868] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1542027916] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542028047] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.17 ms
[1542028071] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542028101] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542028194] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542028195] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.28 ms
[1542028199] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1542028292] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1542028349] HOST ALERT: dek-k2;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 17.20 ms
[1542028641] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.40 ms
[1542028674] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542028678] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542028802] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542028806] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542028822] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542028852] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542028917] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 5.24 ms
[1542028946] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1542029109] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542029173] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542029267] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542029394] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (192.168.56.200)
[1542029500] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542029542] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542029545] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1542029550] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1542029597] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542029680] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542029778] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542029876] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542030056] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542030149] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542030152] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.70 ms
[1542030155] HOST ALERT: mikro56;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 2.71 ms
[1542030284] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542030477] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542030479] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542030610] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542030632] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.38 ms
[1542030659] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542030674] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.09 ms
[1542030758] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1542031000] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542031084] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542031180] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542031229] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542031264] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542031353] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1542031357] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542031357] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542031683] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542031683] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542031815] HOST ALERT: nev6-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1542031857] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542031863] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542031863] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542031941] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1542031957] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542031957] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1542032037] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.04 seconds)
[1542032171] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1542032256] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542032320] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1542032551] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 6.13 ms
[1542032646] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1542032826] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542032876] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542033056] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542033131] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 153.52 ms
[1542033150] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542033476] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542033657] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1542033750] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.12 ms
[1542033756] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542033909] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1542034056] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542034076] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542034120] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1542034205] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 202.22 ms
[1542034213] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.51 ms
[1542034746] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542034831] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542034839] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542034926] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542034973] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542034980] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542034980] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542034986] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.16 ms
[1542035019] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.03 seconds)
[1542035311] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542035491] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542035530] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542035624] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.50 ms
[1542036095] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1542036129] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542036193] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.85 ms
[1542036249] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542036297] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542036304] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1542036304] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1542036308] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1542036634] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542036699] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 8.52 ms
[1542036819] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1542036883] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1542036902] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542036924] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542037028] SERVICE ALERT: mikro56;Ping;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 2.07 ms
[1542037048] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.22 ms
[1542037058] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542037152] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542037156] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 17.86 ms
[1542037209] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542037254] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542037389] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542037434] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.03 seconds)
[1542037482] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542037531] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542037637] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.33 ms
[1542037675] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542037799] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 202.37 ms
[1542037803] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542037993] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542038379] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542038426] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542038429] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542038504] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.38 ms
[1542038550] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542038554] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542038709] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542038754] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1542038880] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542038880] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1542038889] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542038945] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.87 ms
[1542038990] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1542039109] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542039219] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1542039399] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542039456] SERVICE ALERT: dek-k2;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 80%, RTA = 2.21 ms
[1542039702] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542039705] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542039769] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1542039777] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542039826] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 13.06 ms
[1542039902] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.34 ms
[1542040003] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1542040055] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542040107] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542040171] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.91 ms
[1542040453] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542040497] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542040505] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 10.33 ms
[1542040604] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542040621] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1542040747] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1542040754] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542040755] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;2;CRITICAL - Plugin timed out
[1542040818] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542040831] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542040904] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542040918] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1542040968] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542040985] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542041128] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542041135] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542041252] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542041292] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542041294] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542041315] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.02 seconds)
[1542041328] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1542041475] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1542041700] HOST ALERT: mikro31;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542041758] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542041764] HOST ALERT: mikro31;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 16.28 ms
[1542041823] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 2.88 ms
[1542041859] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 14.01 ms
[1542041955] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542042080] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (192.168.56.200)
[1542042204] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542042448] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542042456] SERVICE ALERT: dek-k2;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.37 ms
[1542042486] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 31.02 seconds)
[1542042548] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.23 ms
[1542042560] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542042580] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1542042584] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.52 ms
[1542042624] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542042910] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542042950] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542042974] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1542043079] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542043107] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1542043230] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542043354] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1542043408] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542043409] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1542043589] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542043600] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542043710] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542043728] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542043774] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1542043958] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.82 ms
[1542044054] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542044100] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542044235] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542044255] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.32 ms
[1542044260] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542044384] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1542044387] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542044555] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542044581] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542044601] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542044619] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542044651] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542044705] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542044710] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542044725] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1542044775] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1542044864] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.90 ms
[1542044884] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542044945] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542044948] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1542044956] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1542045100] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 4.16 ms
[1542045190] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1542045274] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542045310] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1542045370] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542045585] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1542045636] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542045725] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542045790] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.58 ms
[1542045791] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.58 ms
[1542045852] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542045859] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542045871] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542045874] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542045884] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542045905] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542045914] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542045914] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1542045929] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542045979] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1542045984] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1542045990] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542045994] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1542045997] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1542046010] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1542046012] HOST ALERT: stach88-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046040] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1542046085] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542046094] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542046100] SERVICE ALERT: metal-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542046103] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542046109] SERVICE ALERT: blago-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542046135] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1542046136] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1542046139] HOST ALERT: blago-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046152] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542046155] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1542046160] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542046179] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1542046184] SERVICE ALERT: sadov-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542046184] SERVICE ALERT: lunach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542046187] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542046210] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1542046210] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1542046214] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046223] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1542046226] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542046229] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1542046229] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1542046232] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1542046244] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1542046253] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1542046265] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1542046269] HOST ALERT: tur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046276] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542046281] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1542046285] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1542046307] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1542046310] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1542046335] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542046336] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1542046352] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1542046381] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1542046394] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542046395] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1542046402] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1542046416] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542046449] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1542046451] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542046481] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.14 ms
[1542046572] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1542046572] HOST ALERT: bkuna-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046577] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1542046619] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046698] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1542046739] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1542046751] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1542046774] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542046799] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542046803] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046824] SERVICE ALERT: varsh-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542046827] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542046889] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1542046895] SERVICE ALERT: eliz-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542046900] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1542046901] HOST ALERT: eliz-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542046923] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542046929] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1542046947] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1542046953] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1542047018] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1542047027] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1542047091] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1542047163] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1542047163] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542047321] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542047343] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.02 seconds)
[1542047615] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542047652] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542047679] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542047709] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 20.89 ms
[1542047752] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542047778] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1542047795] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542047797] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1542047876] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1542047880] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1542048206] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1542048332] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542048335] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542048386] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542048456] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1542048460] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542048481] HOST ALERT: mikro56;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 1.88 ms
[1542048786] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542048941] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1542049060] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.78 ms
[1542049103] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542049106] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542049227] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1542049231] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542049302] HOST FLAPPING ALERT: dek-k2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1542049480] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542049483] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542049508] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542049528] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542049533] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542049592] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1542049603] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1542049609] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1542049659] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1542049661] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1542049663] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542049669] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1542049683] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542049727] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1542049777] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1542049808] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 2.15 ms
[1542049831] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542049918] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542049957] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1542050053] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542050098] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542050141] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1542050167] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542050233] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542050289] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542050291] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1542050292] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542050299] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1542050402] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1542050411] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.77 ms
[1542050412] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1542050418] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1542050420] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542050423] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542050454] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542050543] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1542050546] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1542050578] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1542051034] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542051037] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542051084] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542051158] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1542051162] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1542051183] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.49 ms
[1542051310] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542051322] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542051448] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1542051488] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542051490] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542051595] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1542051668] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542051679] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542051702] SERVICE FLAPPING ALERT: dek-k2;Ping;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1542051762] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.40 ms
[1542051805] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542051808] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542051929] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542051933] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1542052259] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542052365] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1542052385] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542052386] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1542052429] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1542052489] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542052492] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542052509] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1542052520] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542052523] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542052533] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1542052612] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1542052618] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1542052629] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1542052659] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1542052671] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542052703] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542052755] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542052935] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.02 seconds)
[1542052952] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542052953] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542052954] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542053012] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542053039] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542053042] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542053061] SERVICE ALERT: msk197-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542053064] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542053090] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542053103] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542053113] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1542053119] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542053133] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542053134] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542053137] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542053146] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1542053156] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542053162] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1542053165] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1542053167] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 65.05 ms
[1542053182] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542053184] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1542053190] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1542053216] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1542053245] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1542053248] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1542053251] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542053280] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1542053308] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1542053317] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542053318] SERVICE ALERT: nev6-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1542053321] HOST ALERT: nev6-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1542053432] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1542053441] SERVICE ALERT: nev6-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.22)
[1542053444] HOST ALERT: nev6-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.22)
[1542053509] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542053736] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542053740] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1542053763] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542053827] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1542053887] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542053920] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542054153] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542054333] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542054460] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542054464] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542054507] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542054637] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542054790] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542054963] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542054970] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542055064] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1542055143] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542055231] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542055235] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.90 ms
[1542055561] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542055661] SERVICE ALERT: dek-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.12 ms
[1542055690] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542055741] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542055785] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1542055789] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1542055835] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542056132] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542056196] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.88 ms
[1542056415] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542056480] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.61 ms
[1542056522] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542056702] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542057006] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1542057012] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542057058] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542057137] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.76 ms
[1542057183] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 40%, RTA = 2.16 ms
[1542057342] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542057388] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542057522] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542057542] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1542057740] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1542057786] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1542058066] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1542058084] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542058148] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1542058246] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542058340] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 53.50 ms
[1542058446] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542058510] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542058836] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542058965] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1542059291] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542059471] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542059563] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542059566] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542059609] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1542059688] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.91 ms
[1542059692] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1542059734] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 224.62 ms
[1542059738] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1542060018] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542060041] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542060105] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1542060198] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542060291] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1542060431] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542060561] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1542060617] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542060797] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542060887] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542060891] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1542060937] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1542061217] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542061263] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542061397] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542061443] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542061491] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.77 ms
[1542061537] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1542061817] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542061997] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542062092] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.31 ms
[1542062717] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542062763] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542062782] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.42 ms
[1542062828] HOST ALERT: mikro56;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.81 ms
[1542063314] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542063360] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542063465] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1542063483] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (192.168.56.200)
[1542063589] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1542063608] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 60%, RTA = 2.21 ms
[1542063644] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542063690] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542063754] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 527.33 ms
[1542063824] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542064129] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1542064234] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1542064358] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1542064564] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542064744] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1542064962] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1542065288] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542065443] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.36 ms
[1542066255] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1542066410] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.43 ms
[1542066589] SERVICE FLAPPING ALERT: dek-k1;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1542066644] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1542066708] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.22 ms