[1541036196] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541036260] HOST ALERT: mikro70;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541037126] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541037190] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1541042267] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541042530] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.26 ms
[1541042530] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541042534] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.31 ms
[1541042563] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.60 ms
[1541042573] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541042937] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.05 ms
[1541043002] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541043006] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541043032] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 29.39 ms
[1541043036] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.15 ms
[1541043068] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541043072] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541043213] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541043265] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541043352] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1541043356] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541043384] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.75 ms
[1541043445] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1541043494] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.74 ms
[1541043495] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1541043498] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541043498] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1541043507] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541043515] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.85 ms
[1541043519] HOST ALERT: nev6-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541043535] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541043558] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541043559] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541043561] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.45 ms
[1541043589] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1541043618] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541043648] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1541043674] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541043690] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.84 ms
[1541043690] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541043694] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1541043719] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541043723] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1541043738] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541043760] SERVICE ALERT: nev6-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1541043771] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.21 ms
[1541043851] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.26 ms
[1541044491] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541044555] HOST ALERT: kamen-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541044726] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.56 ms
[1541044742] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.79 ms
[1541045051] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541045103] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.26 ms
[1541046165] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.14 ms
[1541046398] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.58 ms
[1541046759] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541046823] HOST ALERT: mikro70;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541047039] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541047041] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541047051] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541047145] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.98 ms
[1541047184] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541047188] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541047353] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541047357] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541047382] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541047394] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541047412] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1541047684] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541047994] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541047998] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.34 ms
[1541048187] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541048252] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 9.23 ms
[1541048284] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.11 ms
[1541048313] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.92 ms
[1541048488] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.98 ms
[1541048514] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.47 ms
[1541048884] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1541048888] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 14.95 ms
[1541049223] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.94 ms
[1541049492] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.50 ms
[1541049493] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.30 ms
[1541049651] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.98 ms
[1541050028] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1541050031] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541050044] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541050121] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541050178] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.23 ms
[1541050424] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.87 ms
[1541050499] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.90 ms
[1541050500] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.79 ms
[1541050504] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.60 ms
[1541050539] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541050691] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.56 ms
[1541050704] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541050791] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.41 ms
[1541050809] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.97 ms
[1541050813] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541050872] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541051120] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.25 ms
[1541051250] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1541051254] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541051319] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.02 ms
[1541051367] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.69 ms
[1541051411] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.83 ms
[1541051415] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.54 ms
[1541051625] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.60 ms
[1541051854] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.28 ms
[1541051864] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.44 ms
[1541052160] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.06 ms
[1541052273] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.73 ms
[1541057861] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1541058180] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541058742] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1541059283] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.24 ms
[1541064852] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541064854] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541064868] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541064878] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541064878] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541064908] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541064915] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541064930] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065020] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065029] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065032] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065036] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065048] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065049] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065060] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065068] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065087] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065088] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065089] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065093] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541065095] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065139] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065200] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065216] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065229] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065240] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065248] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065267] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065269] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065273] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541065310] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065317] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065346] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065357] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065357] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065400] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065407] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541065933] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541066059] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1541066080] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541066663] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541066804] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541067482] SERVICE ALERT: mikro27;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541067606] SERVICE ALERT: mikro27;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541078467] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541078531] HOST ALERT: garkav-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.72 ms
[1541081118] SERVICE ALERT: esx2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541081242] SERVICE ALERT: esx2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1541088208] HOST ALERT: sad29-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541088272] HOST ALERT: sad29-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541088465] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541088480] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541088616] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541088660] HOST ALERT: nevsky130-k1;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541089498] HOST ALERT: sad29-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541089610] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541089678] HOST ALERT: sad29-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541099022] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541099026] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541099033] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541099036] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541099036] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541099036] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541099037] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541099041] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.25 ms
[1541099043] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541099049] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541099051] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541099054] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.44 ms
[1541099054] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541099056] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541099058] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.60 ms
[1541099060] HOST ALERT: garkav-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541099062] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541099064] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541099068] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541099071] SERVICE ALERT: tur-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541099074] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541099079] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541099081] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541099092] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541099098] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541099098] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541099100] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541099104] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541099114] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541099126] SERVICE ALERT: msk197-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541099128] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541099129] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541099132] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541099132] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541099148] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541099151] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541099152] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541099154] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1541099155] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.09 ms
[1541099156] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541099158] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541099159] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541099159] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541099162] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541099162] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541099162] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541099167] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541099175] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541099177] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541099177] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541099182] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541099186] HOST ALERT: garkav-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.21)
[1541099188] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541099190] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541099191] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541099194] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541099194] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541099194] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541099200] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541099204] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541099216] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1541099217] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541099218] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541099220] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541099221] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1541099222] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541099223] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1541099224] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541099225] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541099226] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541099227] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1541099232] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541099239] HOST ALERT: nev6-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541099239] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541099242] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541099243] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541099249] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541099249] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541099250] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541099252] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541099259] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541099268] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541099271] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541099279] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541099294] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541099296] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541099307] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541099338] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541099345] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541099346] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541099351] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541099362] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541099365] HOST ALERT: nev6-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.21)
[1541099365] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541099369] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541099375] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541099385] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541099392] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541099394] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541099405] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541099419] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541099422] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541099424] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541099427] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541099433] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541099438] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541099441] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541099449] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541099460] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541099475] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541099480] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.73 ms
[1541099522] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541099526] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541099534] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541099547] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541099553] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541099560] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541099564] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541099570] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541099589] SERVICE ALERT: nev6-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.21)
[1541099596] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541099606] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541099662] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541099847] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541099890] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541100027] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541101477] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541101603] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541101656] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541101760] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541101782] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541101855] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541102640] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541102643] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541102681] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541102763] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541102769] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541102785] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541102786] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541102807] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541102807] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541102837] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541102933] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541102958] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541102961] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541102965] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541103080] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541103081] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541103084] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541103110] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541103111] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541103118] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541103141] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541103154] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541103231] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541103262] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541103280] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541103290] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541103322] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 31.01 seconds)
[1541103396] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541103399] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541103434] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541103442] HOST ALERT: nev6-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.39.22)
[1541103453] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541103502] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541103519] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541103525] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541103533] SERVICE ALERT: nev6-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.39.22)
[1541103568] HOST ALERT: nev6-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.39.22)
[1541103633] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541109075] SERVICE ALERT: mikro72;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541109199] SERVICE ALERT: mikro72;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.85 ms