[1541463427] SERVICE ALERT: mikro54;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.200)
[1541463551] SERVICE ALERT: mikro54;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.83 ms
[1541474763] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.85 ms
[1541474850] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541474854] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541474860] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541474864] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541474879] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541474960] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1541474964] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1541475004] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1541475008] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.72 ms
[1541475051] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.79 ms
[1541475124] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541475135] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541475137] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541475139] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541475342] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.15 ms
[1541475379] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.87 ms
[1541475383] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541475408] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541475514] SERVICE ALERT: nev-k9;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541475517] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541475637] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541475665] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541475697] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541475699] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.73 ms
[1541475847] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1541475863] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.84 ms
[1541475867] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.53 ms
[1541475927] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541475990] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541475994] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.31 ms
[1541476012] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541476016] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.78 ms
[1541476052] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.88 ms
[1541476056] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541476060] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541476122] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.84 ms
[1541476265] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541476269] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541476273] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541476277] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541476326] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1541476330] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541476352] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.69 ms
[1541476407] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541476420] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1541476490] SERVICE ALERT: nev-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541476493] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541476594] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1541476633] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.82 ms
[1541476637] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.14 ms
[1541476641] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541476673] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541476830] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1541477877] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1541478210] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1541478214] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541478332] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541478414] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1541478922] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541479144] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.37 ms
[1541479240] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541479244] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541479298] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.13 ms
[1541479341] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541479493] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.77 ms
[1541479568] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.91 ms
[1541479723] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1541479747] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1541480005] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.18 ms
[1541480737] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1541480748] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.17 ms
[1541480752] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 23.26 ms
[1541480848] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1541480852] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541481077] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 6.91 ms
[1541481341] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 612.39 ms
[1541481384] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541481437] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.90 ms
[1541481441] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541481635] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 11.55 ms
[1541481803] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541481867] HOST ALERT: nev-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541481879] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1541481946] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.17 ms
[1541482085] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.43 ms
[1541482224] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.55 ms
[1541482255] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.65 ms
[1541482264] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.21 ms
[1541482292] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.40 ms
[1541482296] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.38 ms
[1541482379] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.70 ms
[1541482454] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.11 ms
[1541482458] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.93 ms
[1541482514] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.50 ms
[1541482558] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1541482986] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.95 ms
[1541482990] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.44 ms
[1541483210] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.44 ms
[1541483214] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.39 ms
[1541483334] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541483339] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1541483412] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.98 ms
[1541483577] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1541483712] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.38 ms
[1541483716] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.06 ms
[1541483720] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.53 ms
[1541483725] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.75 ms
[1541483835] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1541483949] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.50 ms
[1541484198] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1541484312] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541484321] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1541484487] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 15.94 ms
[1541484499] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541484514] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541484518] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1541484777] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.68 ms
[1541487410] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.71 ms
[1541487620] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.74 ms
[1541487701] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.97 ms
[1541488238] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541488575] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541488699] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.96 ms
[1541489738] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.29 ms
[1541489742] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541491340] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541491349] SERVICE ALERT: mikro63;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541491404] HOST ALERT: engels51-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.31 ms
[1541491473] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.30 ms
[1541493665] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541493729] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541495048] HOST ALERT: 5sovet-tel1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541495112] HOST ALERT: 5sovet-tel1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 13.61 ms
[1541495438] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541495502] HOST ALERT: garkav-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.47 ms
[1541498112] SERVICE ALERT: mikro70;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541498236] SERVICE ALERT: mikro70;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1541505194] SERVICE ALERT: mikro31;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.31.200)
[1541505194] HOST ALERT: mikro31;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.200)
[1541505195] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541505201] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541505201] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541505206] SERVICE ALERT: sad29-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541505206] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541505206] HOST ALERT: mikro58;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.200)
[1541505211] HOST ALERT: varsh-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505213] HOST ALERT: stach88-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505213] SERVICE ALERT: lens22-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541505215] HOST ALERT: zahar-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.36.21)
[1541505216] HOST ALERT: lens22-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505216] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541505216] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541505216] SERVICE ALERT: mikro47;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.200)
[1541505216] HOST ALERT: mikro47;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.200)
[1541505222] HOST ALERT: mikro42;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.42.200)
[1541505229] HOST ALERT: turist28-k1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541505230] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505234] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541505239] HOST ALERT: mikro60;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505242] SERVICE ALERT: kirochn-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.4.21)
[1541505244] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505245] HOST ALERT: kirochn-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.21)
[1541505249] SERVICE ALERT: 5sovet-tel1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541505258] HOST ALERT: mikro31;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 17.02 ms
[1541505259] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541505265] HOST ALERT: suvor42-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1541505270] HOST ALERT: sad29-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.19 ms
[1541505270] HOST ALERT: mikro58;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541505275] HOST ALERT: varsh-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.54 ms
[1541505277] HOST ALERT: stach88-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 10.99 ms
[1541505279] HOST ALERT: zahar-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.77 ms
[1541505280] HOST ALERT: lens22-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.85 ms
[1541505280] HOST ALERT: nevsk104-k4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1541505280] HOST ALERT: mikro47;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541505286] HOST ALERT: mikro42;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.06 ms
[1541505289] HOST ALERT: kirochn-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.22)
[1541505293] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.43 ms
[1541505294] HOST ALERT: engels51-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.82 ms
[1541505303] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541505308] HOST ALERT: nev-k7;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541505309] HOST ALERT: kirochn-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 57.27 ms
[1541505318] SERVICE ALERT: mikro31;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.92 ms
[1541505325] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541505330] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541505337] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541505340] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1541505340] SERVICE ALERT: mikro47;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541505353] HOST ALERT: kirochn-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 55.01 ms
[1541505358] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541505366] SERVICE ALERT: kirochn-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 74.61 ms
[1541505373] SERVICE ALERT: 5sovet-tel1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.68 ms
[1541505428] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541505553] HOST ALERT: garkav-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.54 ms
[1541507192] HOST ALERT: mikro68;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.200)
[1541507251] HOST ALERT: mikro68;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541511561] HOST ALERT: mikro36;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.36.200)
[1541511622] HOST ALERT: mikro36;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.04 ms
[1541512306] SERVICE ALERT: mikro46;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.46.200)
[1541512430] SERVICE ALERT: mikro46;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.95 ms
[1541513045] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541513169] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.94 ms
[1541513579] HOST ALERT: mikro63;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541513643] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.39 ms
[1541513658] SERVICE ALERT: mikro25;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.200)
[1541513782] SERVICE ALERT: mikro25;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.76 ms
[1541515828] HOST ALERT: mikro40;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541515892] HOST ALERT: mikro40;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1541516128] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541516192] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.04 ms
[1541520089] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541520123] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541520215] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541520417] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541520420] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541520568] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541520600] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541521440] SERVICE ALERT: mikro39;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.39.200)
[1541521521] SERVICE ALERT: sad29-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541521524] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541521564] SERVICE ALERT: mikro39;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541521644] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541521647] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541523967] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541523970] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541524090] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541524093] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541527247] SERVICE ALERT: nev-k11;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541527371] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541527377] SERVICE ALERT: stach88-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541527380] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541527380] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541527479] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541527500] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541527506] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541527506] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541527511] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541527513] HOST ALERT: dek-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541527520] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541527524] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541527557] SERVICE ALERT: engels-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541527560] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541527568] HOST ALERT: lunach-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541527596] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541527632] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541527634] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541527635] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541527637] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541527639] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541527659] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541527680] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541527686] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541527697] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541527721] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541527741] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541527743] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541527748] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541527755] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541527761] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541527761] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541527764] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541527779] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541527821] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541527834] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541527845] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541527854] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541527867] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541527869] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541527878] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541527890] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541527891] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541527892] SERVICE ALERT: sadov-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541527894] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541527895] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541527905] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541527947] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541527971] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541527979] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541528004] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541528014] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541528020] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541528043] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541528051] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541528054] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541528075] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541528092] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541528093] HOST ALERT: engels51-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541528105] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541528113] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541528124] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541528136] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541528174] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541528180] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541528203] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541528208] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541528219] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541528247] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541528299] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541528323] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541528350] SERVICE ALERT: mikro54;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.54.200)
[1541528383] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541528474] SERVICE ALERT: mikro54;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.96 ms
[1541528500] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541528578] SERVICE ALERT: varsh-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541528581] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541528701] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541528707] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541528733] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541528859] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541529075] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541529315] SERVICE ALERT: tur-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541529318] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541529438] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541529438] HOST ALERT: msk197-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541529441] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541529564] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541529734] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541529736] SERVICE ALERT: kamen-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541529739] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541529859] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541529865] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541530277] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541530398] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541530457] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.00 seconds)
[1541530898] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541530901] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541531021] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541531027] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541531079] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541531096] HOST ALERT: sad29-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541531129] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541531139] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541531165] SERVICE ALERT: nev-k5;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541531168] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541531182] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541531202] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541531221] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541531237] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541531240] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541531308] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541531316] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541531348] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541531360] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541531366] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541531515] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541531641] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541531678] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541531743] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541531813] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541531867] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541532039] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541533263] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541533443] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541533474] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541534289] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541534415] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541534431] SERVICE ALERT: kamen-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541534434] HOST ALERT: kamen-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541534481] HOST ALERT: msk197-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541534551] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541534554] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541534557] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541534607] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541534734] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541534741] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541534778] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541534781] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541534828] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541534847] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541534850] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541534859] SERVICE ALERT: kamen-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541534862] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541534901] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541534904] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541534914] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541534941] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541534970] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541534976] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541534982] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541534985] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541535008] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541535064] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541535155] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541535281] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541535385] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541535415] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541535451] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541535536] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541535595] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541535893] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541536021] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541536041] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541536073] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541536201] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541536257] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out