[1541555124] SERVICE ALERT: mikro72;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541555127] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541555191] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.02 ms
[1541555248] SERVICE ALERT: mikro72;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.83 ms
[1541558774] SERVICE ALERT: mikro41;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.200)
[1541558898] SERVICE ALERT: mikro41;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1541559858] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.73 ms
[1541559862] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541560065] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541560069] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.81 ms
[1541561018] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1541561047] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1541561112] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541561198] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1541561465] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.51 ms
[1541561555] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541561559] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1541561588] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.60 ms
[1541561660] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541561664] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.65 ms
[1541561689] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541561909] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1541561913] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.26 ms
[1541561919] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541561965] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.87 ms
[1541561970] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.95 ms
[1541561971] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541561975] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541561983] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1541562004] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541562005] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1541562181] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541562208] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1541562223] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 1.66 ms
[1541562227] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541562247] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541562502] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541562542] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541562706] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.86 ms
[1541562846] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.75 ms
[1541562971] HOST ALERT: blago-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.19 ms
[1541563063] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1541563137] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541563285] SERVICE ALERT: blago-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.08 ms
[1541563863] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541563889] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541563927] HOST ALERT: mikro69;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.69 ms
[1541563953] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.74 ms
[1541564302] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.80 ms
[1541564558] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.51 ms
[1541564979] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1541564987] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1541565216] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.41 ms
[1541565224] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541565694] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1541566045] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1541566048] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541566049] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1541566053] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.93 ms
[1541566057] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541566071] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.27 ms
[1541566091] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541566095] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1541566363] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541567366] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 10.30 ms
[1541567523] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.20 ms
[1541567575] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.72 ms
[1541567578] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.80 ms
[1541567933] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.36 ms
[1541567956] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.79 ms
[1541568012] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541568199] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541568203] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.47 ms
[1541568225] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.74 ms
[1541568302] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541568373] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1541568416] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.45 ms
[1541568647] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.76 ms
[1541568792] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.97 ms
[1541568811] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1541568870] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.89 ms
[1541568910] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.61 ms
[1541568911] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 2.47 ms
[1541569002] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541569037] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.55 ms
[1541569039] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.37 ms
[1541569041] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.66 ms
[1541569043] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.71 ms
[1541569081] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.80 ms
[1541569124] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.47 ms
[1541569158] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541569161] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.78 ms
[1541569171] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.51 ms
[1541569309] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.09 ms
[1541569415] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.46 ms
[1541569575] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.23 ms
[1541569579] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.84 ms
[1541569641] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.26 ms
[1541569643] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.26 ms
[1541569643] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541569648] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541569660] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.59 ms
[1541569708] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.35 ms
[1541569810] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.10 ms
[1541569823] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 18.58 ms
[1541570101] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 11.87 ms
[1541570289] SERVICE ALERT: nev-k5;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.93 ms
[1541570293] HOST ALERT: nev-k5;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1541570361] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.59 ms
[1541570472] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.67 ms
[1541570476] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.39 ms
[1541570676] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541570685] HOST ALERT: nevsk104-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1541570722] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.05 ms
[1541570740] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1541570975] SERVICE ALERT: nevsk104-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.85 ms
[1541570992] HOST ALERT: turist28-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.21)
[1541571056] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.79 ms
[1541571105] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541571169] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1541572882] HOST ALERT: turist28-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541572946] HOST FLAPPING ALERT: turist28-k1;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1541572946] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1541573249] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541573313] HOST ALERT: engels51-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.22 ms
[1541573595] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541573601] HOST ALERT: metal-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541573605] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541573659] HOST FLAPPING ALERT: prosv48-k1;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1541573659] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 7.68 ms
[1541573665] HOST ALERT: metal-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.67 ms
[1541573729] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.81 ms
[1541574872] SERVICE ALERT: mikro63;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541574875] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541574885] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541574891] HOST ALERT: metal-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541574894] SERVICE ALERT: turist28-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.72.21)
[1541574897] HOST ALERT: turist28-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.21)
[1541574931] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541574933] HOST FLAPPING ALERT: metal-k1;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1541574933] HOST ALERT: metal-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.70 ms
[1541574939] HOST ALERT: mikro63;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.02 ms
[1541574949] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 9.85 ms
[1541574961] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.88 ms
[1541574996] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.17 ms
[1541575018] SERVICE ALERT: turist28-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.07 ms
[1541575111] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541575379] HOST ALERT: mikro72;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541575408] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541575443] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.83 ms
[1541575538] HOST ALERT: mikro63;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541575556] SERVICE ALERT: metal-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541575559] HOST ALERT: metal-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541575575] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541575587] HOST ALERT: turist28-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541575595] SERVICE ALERT: mikro63;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.200)
[1541575602] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541575616] SERVICE ALERT: turist28-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.72.21)
[1541575619] SERVICE ALERT: mikro69;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541575622] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541575647] SERVICE ALERT: mikro72;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541575650] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541575664] HOST ALERT: mikro63;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.200)
[1541575680] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.93 ms
[1541575684] HOST ALERT: metal-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.79 ms
[1541575686] HOST ALERT: mikro69;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.84 ms
[1541575702] HOST ALERT: prosv48-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1541575714] HOST ALERT: turist28-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1541575714] HOST FLAPPING ALERT: mikro72;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1541575714] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.32 ms
[1541575719] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.98 ms
[1541575723] HOST FLAPPING ALERT: mikro63;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1541575723] HOST ALERT: mikro63;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.26 ms
[1541575726] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1541575740] SERVICE ALERT: turist28-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.38 ms
[1541575743] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.72 ms
[1541575771] SERVICE ALERT: mikro72;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1541576075] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541576079] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.65 ms
[1541576403] HOST ALERT: mikro60;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.60.200)
[1541576467] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.03 ms
[1541576615] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1541576732] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.99 ms
[1541577610] SERVICE ALERT: kamen-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541577615] HOST ALERT: kamen-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541577679] HOST ALERT: kamen-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1541577734] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.02 ms
[1541577978] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541578030] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541578089] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541578104] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541578332] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541578708] HOST FLAPPING ALERT: kamen-k3;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1541578708] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.12 ms
[1541578723] HOST FLAPPING ALERT: mikro63;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541579002] HOST FLAPPING ALERT: prosv48-k1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1541579014] HOST FLAPPING ALERT: turist28-k1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1541579056] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541579280] HOST FLAPPING ALERT: metal-k1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1541579371] HOST FLAPPING ALERT: mikro72;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541579569] HOST ALERT: tur-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541579723] HOST ALERT: tur-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.11 ms
[1541581456] HOST FLAPPING ALERT: kamen-k3;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1541581678] HOST ALERT: mikro68;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541581742] HOST ALERT: mikro68;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1541592544] HOST ALERT: mikro68;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.200)
[1541592608] HOST ALERT: mikro68;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541595323] HOST ALERT: mikro36;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.36.200)
[1541595387] HOST ALERT: mikro36;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1541597416] HOST ALERT: mikro70;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541597480] HOST ALERT: mikro70;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541600218] SERVICE ALERT: mikro57;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541600342] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.65 ms
[1541602143] SERVICE ALERT: nev-k2;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 2.68 ms
[1541602267] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1541603898] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541603901] HOST ALERT: dek-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541604021] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.21)
[1541604024] HOST ALERT: dek-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.21)
[1541606318] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541606444] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541606646] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541606695] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541606759] HOST ALERT: nev-k3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1541607437] SERVICE ALERT: mikro70;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.70.200)
[1541607561] SERVICE ALERT: mikro70;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.31 ms
[1541608647] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541608647] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541608770] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541609074] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541609077] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541609197] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541609200] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541613757] SERVICE ALERT: lunach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541613760] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541613767] HOST ALERT: grivcov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541613770] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541613814] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541613817] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541613880] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541613886] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541613893] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541613896] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541613937] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541613943] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541613970] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541613993] HOST ALERT: suvor42-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541614002] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541614042] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541614047] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541614047] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541614071] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541614080] SERVICE ALERT: metal-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541614080] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541614096] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541614101] HOST ALERT: prosv48-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541614102] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541614119] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541614125] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541614156] HOST ALERT: garkav-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.45.21)
[1541614156] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541614156] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541614159] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541614159] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541614197] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541614201] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541614204] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541614205] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541614209] HOST ALERT: stach88-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541614222] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541614227] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541614228] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541614271] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.45.21)
[1541614279] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541614279] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541614284] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541614285] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541614285] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541614335] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541614339] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541614348] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541614362] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541614385] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541614398] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541614418] HOST ALERT: engels-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541614423] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541614443] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541614465] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541614500] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541614509] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541614511] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541614544] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541614549] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541614572] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541614608] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541614615] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541614635] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541614636] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541614637] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541614638] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541614660] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541614680] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541614698] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541614734] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541614741] SERVICE ALERT: mikro57;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541614746] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541614770] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541614774] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541614810] HOST ALERT: mikro57;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.15 ms
[1541614811] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541614847] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541614865] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.69 ms
[1541614937] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541615051] SERVICE ALERT: sadov-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541615081] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541615101] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541615202] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541615261] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541615332] SERVICE ALERT: engels51-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541615339] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541615455] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541615465] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541615493] HOST ALERT: mikro60;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541615496] SERVICE ALERT: mikro60;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541615557] HOST ALERT: mikro60;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541615620] SERVICE ALERT: mikro60;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.90 ms
[1541615685] SERVICE ALERT: msk197-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541615688] HOST ALERT: msk197-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541615808] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541615814] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541615963] HOST ALERT: kamen-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541616089] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541616259] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541616534] HOST ALERT: bkuna-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541616622] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541616659] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541616821] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541616824] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541616944] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541616947] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541617652] SERVICE ALERT: sad29-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541617655] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541617761] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541617764] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541617775] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541617781] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541617787] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541617884] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541617890] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541617932] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541617967] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541617994] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541618011] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541618058] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541618119] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541618137] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541618149] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541618201] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541618262] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541618290] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541618619] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541618799] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541618916] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541619574] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541619694] SERVICE ALERT: nev-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541619716] HOST ALERT: mikro43;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.200)
[1541619724] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541619754] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541619780] HOST ALERT: mikro43;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1541619845] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541619904] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541620025] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541620538] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541620718] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541620892] SERVICE ALERT: nev-k7;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541620922] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541620954] SERVICE ALERT: kamen-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541620954] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541620998] HOST ALERT: msk197-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541621007] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541621015] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541621043] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541621072] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541621077] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541621102] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541621124] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541621133] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541621198] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)
[1541621264] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541621301] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541621307] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541621320] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541621336] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541621337] SERVICE ALERT: nev-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541621339] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541621370] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541621373] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541621433] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541621488] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541621493] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541621499] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541621519] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541621655] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541621785] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541621942] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541621965] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541621992] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541622122] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541622359] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541624465] SERVICE ALERT: mikro57;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541624589] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.23 ms
[1541628396] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541628520] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1541631034] HOST ALERT: mikro04;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541631035] HOST ALERT: kirochn-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541631054] HOST ALERT: kirochn-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.22)
[1541631070] SERVICE ALERT: mikro04;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.200)
[1541631160] HOST ALERT: mikro04;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.200)
[1541631180] HOST ALERT: kirochn-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.22)
[1541631188] HOST ALERT: kirochn-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.4.21)
[1541631365] SERVICE ALERT: kirochn-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.21)
[1541631574] SERVICE ALERT: kirochn-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.4.22)
[1541632589] HOST ALERT: mikro48;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541632595] HOST ALERT: mikro52;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541632595] HOST ALERT: mikro57;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541632626] SERVICE ALERT: mikro52;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.52.200)
[1541632684] SERVICE ALERT: mikro61;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541632685] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541632687] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541632715] HOST ALERT: mikro48;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.200)
[1541632721] HOST ALERT: mikro52;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.52.200)
[1541632721] HOST ALERT: mikro57;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.200)
[1541632739] HOST ALERT: dach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.52.21)
[1541632742] SERVICE ALERT: mikro69;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541632807] SERVICE ALERT: mikro61;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541632811] HOST ALERT: mikro69;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.200)
[1541632813] HOST ALERT: mikro61;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.200)
[1541632865] HOST ALERT: dach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.52.21)
[1541632871] SERVICE ALERT: dach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.52.21)
[1541632988] SERVICE ALERT: mikro57;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541633142] SERVICE ALERT: mikro48;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541633319] HOST ALERT: mikro48;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1541633325] HOST ALERT: mikro52;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541633325] HOST ALERT: mikro57;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1541633350] SERVICE ALERT: mikro52;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541633595] SERVICE ALERT: dach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 7.30 ms
[1541633599] HOST ALERT: dach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.25 ms
[1541633712] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.45 ms
[1541633715] HOST ALERT: mikro69;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.65 ms
[1541633717] HOST ALERT: mikro61;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1541633866] SERVICE ALERT: mikro48;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541634008] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541634067] SERVICE ALERT: mikro69;Ping;OK;HARD;1;PING OK - Packet loss = 20%, RTA = 5.40 ms
[1541634311] SERVICE ALERT: mikro57;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541634314] HOST ALERT: mikro69;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541634314] HOST ALERT: mikro57;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541634316] HOST ALERT: mikro61;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541634434] SERVICE ALERT: mikro57;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.200)
[1541634440] HOST ALERT: mikro69;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.200)
[1541634440] HOST ALERT: mikro57;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.200)
[1541634442] HOST ALERT: mikro61;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.200)
[1541634465] SERVICE ALERT: mikro48;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541634468] HOST ALERT: mikro48;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541634498] HOST ALERT: dach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.52.21)
[1541634524] HOST ALERT: mikro52;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.52.200)
[1541634549] SERVICE ALERT: mikro52;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.52.200)
[1541634588] SERVICE ALERT: mikro48;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.200)
[1541634594] HOST ALERT: mikro48;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.200)
[1541634607] SERVICE ALERT: mikro61;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.200)
[1541634624] HOST ALERT: dach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.52.21)
[1541634650] HOST ALERT: mikro52;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.52.200)
[1541634665] SERVICE ALERT: mikro69;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.200)
[1541634731] SERVICE FLAPPING ALERT: mikro61;Ping;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1541634731] SERVICE ALERT: mikro61;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1541634735] HOST FLAPPING ALERT: mikro61;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1541634735] HOST ALERT: mikro61;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1541634799] HOST FLAPPING ALERT: dach-k1;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1541634799] HOST ALERT: dach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.96 ms
[1541634898] HOST FLAPPING ALERT: mikro48;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1541634898] HOST ALERT: mikro48;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541634954] HOST FLAPPING ALERT: mikro52;STARTED; Host appears to have started flapping (20.7% change > 20.0% threshold)
[1541634954] HOST ALERT: mikro52;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541635035] SERVICE FLAPPING ALERT: mikro57;Ping;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1541635035] SERVICE ALERT: mikro57;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.54 ms
[1541635039] HOST FLAPPING ALERT: mikro57;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1541635039] HOST ALERT: mikro57;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.00 ms
[1541635044] HOST FLAPPING ALERT: mikro69;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1541635044] HOST ALERT: mikro69;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.63 ms
[1541635189] SERVICE FLAPPING ALERT: mikro48;Ping;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1541635189] SERVICE ALERT: mikro48;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms