[1541809010] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541809074] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.55 ms
[1541810815] SERVICE FLAPPING ALERT: mikro56;Ping;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541820051] HOST ALERT: nevsky130-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541820096] HOST ALERT: nevsky130-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541820345] SERVICE ALERT: nevsky130-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541820387] SERVICE ALERT: nevsky130-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1541820538] HOST ALERT: nev6-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.31 ms
[1541820541] HOST ALERT: nev6-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1541820565] SERVICE ALERT: mikro63;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 7.88 ms
[1541820568] HOST ALERT: nev-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541820569] SERVICE ALERT: nev6-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1541820576] HOST ALERT: nev-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1541820576] SERVICE ALERT: nev-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1541820577] HOST ALERT: nev-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1541820579] HOST ALERT: nev-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1541820585] SERVICE ALERT: nev-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1541820643] SERVICE ALERT: nev6-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541820689] SERVICE ALERT: mikro63;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.15 ms
[1541820829] SERVICE ALERT: nev-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541820872] SERVICE ALERT: nev-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1541820895] SERVICE ALERT: turist28-k1;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 3.74 ms
[1541821019] SERVICE ALERT: turist28-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.84 ms
[1541821077] HOST ALERT: msk197-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1541821077] HOST ALERT: msk197-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.85 ms
[1541821187] SERVICE ALERT: msk197-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1541821351] SERVICE ALERT: msk197-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541821411] HOST ALERT: nevsky78-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1541821464] HOST ALERT: nevsky78-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.69 ms
[1541821550] SERVICE ALERT: nevsky78-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1541821594] SERVICE ALERT: nevsk104-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.49 ms
[1541821598] HOST ALERT: nevsk104-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.84 ms
[1541821701] SERVICE ALERT: nevsky78-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.57 ms
[1541821717] HOST ALERT: nevsk104-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.67 ms
[1541822002] SERVICE ALERT: nevsk104-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1541822073] HOST ALERT: nev-k11;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.69 ms
[1541822084] HOST ALERT: nev-k6;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1541822090] HOST ALERT: nev-k9;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1541822148] SERVICE ALERT: nev-k6;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1541822194] SERVICE ALERT: nev-k11;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.11 ms
[1541822281] HOST ALERT: nev-k7;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.60 ms
[1541822391] SERVICE ALERT: nev-k9;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1541822522] SERVICE ALERT: nev-k7;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.62 ms
[1541823021] HOST ALERT: ushinsk-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.79 ms
[1541823045] SERVICE ALERT: ushinsk-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.95 ms
[1541823187] SERVICE ALERT: sadov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1541823191] HOST ALERT: sadov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1541823223] SERVICE ALERT: sadov-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1541823227] HOST ALERT: sadov-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.49 ms
[1541823254] SERVICE ALERT: sadov-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.39 ms
[1541823258] HOST ALERT: sadov-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1541823552] HOST ALERT: sad29-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1541823553] HOST ALERT: sad29-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1541823824] SERVICE ALERT: sad29-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1541824045] SERVICE ALERT: sad29-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541824360] SERVICE ALERT: 8line-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1541824364] HOST ALERT: 8line-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541824540] SERVICE ALERT: lunach-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 10.13 ms
[1541824544] HOST ALERT: lunach-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 9.40 ms
[1541824549] HOST ALERT: kamen-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541824552] HOST ALERT: kamen-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541824553] HOST ALERT: kamen-k3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1541824602] SERVICE ALERT: kamen-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.70 ms
[1541824699] SERVICE ALERT: kamen-k3;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.55 ms
[1541824707] SERVICE ALERT: kamen-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541824796] HOST ALERT: 8line-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1541825044] SERVICE ALERT: 8line-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541825076] HOST ALERT: eliz-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541825097] SERVICE ALERT: eliz-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.38 ms
[1541825101] HOST ALERT: eliz-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.35 ms
[1541825367] SERVICE ALERT: eliz-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.33 ms
[1541825668] HOST ALERT: suvor42-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541825844] SERVICE ALERT: suvor42-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.17 ms
[1541825904] HOST ALERT: msk197-k2;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1541825968] HOST ALERT: msk197-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.79 ms
[1541826002] HOST ALERT: bkuna-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 7.68 ms
[1541826187] SERVICE ALERT: prosv87-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.96 ms
[1541826191] HOST ALERT: prosv87-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1541826240] HOST ALERT: prosv87-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541826306] SERVICE ALERT: bkuna-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.96 ms
[1541826567] SERVICE ALERT: prosv87-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.66 ms
[1541826927] SERVICE ALERT: engels51-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.09 ms
[1541826927] HOST ALERT: engels51-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 5.40 ms
[1541826963] HOST ALERT: gavan-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.83 ms
[1541827206] HOST ALERT: bryanc20-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.26 ms
[1541827228] SERVICE ALERT: gavan-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.41 ms
[1541827616] SERVICE ALERT: bryanc20-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.43 ms
[1541827779] HOST ALERT: varsh-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.56 ms
[1541828151] SERVICE ALERT: varsh-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.37 ms
[1541828368] HOST ALERT: engels-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.07 ms
[1541828638] HOST ALERT: prosv48-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.48 ms
[1541828649] SERVICE ALERT: metal-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.36 ms
[1541828653] HOST ALERT: metal-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.28 ms
[1541828661] SERVICE ALERT: engels-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.80 ms
[1541828680] HOST ALERT: lens22-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.67 ms
[1541828694] HOST ALERT: stach88-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 11.93 ms
[1541828734] HOST ALERT: gasheka-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.30 ms
[1541828744] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.27 ms
[1541828785] SERVICE ALERT: stach88-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 42.42 ms
[1541828808] SERVICE ALERT: prosv48-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.42 ms
[1541828971] SERVICE ALERT: lens22-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1541828972] HOST ALERT: tur-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.75 ms
[1541829027] SERVICE ALERT: gasheka-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1541829040] HOST ALERT: grivcov-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.55 ms
[1541829099] SERVICE ALERT: grivcov-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.43 ms
[1541829113] HOST ALERT: garkav-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.52 ms
[1541829219] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.20 ms
[1541829221] HOST ALERT: aviakonstr-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.77 ms
[1541829294] SERVICE ALERT: tur-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.10 ms
[1541829300] HOST ALERT: aviakonstr-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.50 ms
[1541829330] SERVICE ALERT: aviakonstr-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.52 ms
[1541829512] SERVICE ALERT: aviakonstr-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.41 ms
[1541829558] SERVICE ALERT: garkav-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.56 ms
[1541829947] SERVICE ALERT: nevsk104-k4;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.81 ms
[1541829951] HOST ALERT: nevsk104-k4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.87 ms
[1541829954] HOST ALERT: mikro61;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541830018] HOST ALERT: mikro61;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 8.82 ms
[1541851708] HOST ALERT: dach-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541851772] HOST ALERT: dach-k1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.80 ms
[1541855742] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 60%, RTA = 2.33 ms
[1541855885] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541855950] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.98 ms
[1541856036] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541856187] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;2;CRITICAL - Plugin timed out
[1541856311] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1541856340] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541856466] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541856966] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541856970] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541857016] SERVICE ALERT: mikro56;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.41 ms
[1541857020] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1541857031] HOST FLAPPING ALERT: dek-k1;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1541857031] HOST ALERT: dek-k1;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.51 ms
[1541857150] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541857346] HOST FLAPPING ALERT: mikro56;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1541857346] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1541857424] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.39 ms
[1541857473] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1541857799] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541857864] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.35 ms
[1541857956] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541858046] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541858136] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541858138] SERVICE ALERT: dek-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541858170] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.23 ms
[1541858262] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.78 ms
[1541858266] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.16 ms
[1541858592] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541858676] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541858741] HOST FLAPPING ALERT: dek-k2;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1541858741] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 4.16 ms
[1541858772] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541858866] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.13 ms
[1541859089] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1541859155] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 2.64 ms
[1541859192] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541859366] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541859372] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541859374] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 10.46 ms
[1541859466] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 19.57 ms
[1541859700] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541859792] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541859880] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541859972] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541859974] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.24 ms
[1541860051] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541860066] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.46 ms
[1541860115] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 18.16 ms
[1541860300] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541860392] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541860441] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541860480] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541860572] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541860574] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 13.49 ms
[1541860621] HOST ALERT: mikro56;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541860666] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.62 ms
[1541860900] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541860992] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541861080] HOST ALERT: dek-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541861172] HOST ALERT: dek-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541861174] HOST ALERT: dek-k2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.30 ms
[1541861220] HOST ALERT: mikro56;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1541861266] HOST ALERT: dek-k1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.10 ms
[1541861500] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541861546] HOST ALERT: mikro56;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541861592] HOST ALERT: dek-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541861655] HOST ALERT: dek-k2;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 2.54 ms
[1541861657] HOST ALERT: dek-k1;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 35.58 ms
[1541861700] HOST ALERT: mikro56;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 22.51 ms
[1541861797] SERVICE ALERT: dek-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541861843] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541861889] SERVICE ALERT: dek-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541861921] SERVICE ALERT: dek-k2;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1541861964] SERVICE ALERT: mikro56;Ping;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (192.168.56.200)
[1541862013] SERVICE ALERT: dek-k1;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.29 ms
[1541862087] SERVICE ALERT: mikro56;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541862146] HOST ALERT: mikro56;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.56.200)
[1541862272] HOST ALERT: mikro56;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.200)
[1541864921] HOST FLAPPING ALERT: dek-k2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1541864975] HOST FLAPPING ALERT: mikro56;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1541865078] HOST ALERT: nevsky130-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541865193] HOST FLAPPING ALERT: dek-k1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1541865204] HOST ALERT: nevsky130-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.21)
[1541865344] SERVICE ALERT: nevsky130-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.21)
[1541865774] HOST ALERT: mikro72;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.72.200)
[1541865838] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.39 ms
[1541868947] HOST ALERT: kirochn-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541869011] HOST ALERT: kirochn-k2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 79.26 ms
[1541869072] SERVICE ALERT: sad29-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541869075] HOST ALERT: sad29-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541869195] SERVICE ALERT: sad29-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.21)
[1541869201] HOST ALERT: sad29-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.21)
[1541871218] HOST ALERT: sadov-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541871219] SERVICE ALERT: sadov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541871398] HOST ALERT: sadov-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541871511] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541871514] HOST ALERT: aviakonstr-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541871599] HOST ALERT: aviakonstr-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541871634] SERVICE ALERT: aviakonstr-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.21)
[1541871637] HOST ALERT: aviakonstr-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.21)
[1541871725] HOST ALERT: aviakonstr-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.24.22)
[1541871759] SERVICE ALERT: 8line-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541871762] HOST ALERT: 8line-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541871882] SERVICE ALERT: 8line-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.21)
[1541871885] HOST ALERT: 8line-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.21)
[1541871929] SERVICE ALERT: aviakonstr-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.24.22)
[1541872926] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541872927] HOST ALERT: dek-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541872929] HOST ALERT: grivcov-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541872939] HOST ALERT: garkav-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541872970] HOST ALERT: blago-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541873031] SERVICE ALERT: blago-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.71.21)
[1541873049] SERVICE ALERT: grivcov-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.27.21)
[1541873052] HOST ALERT: grivcov-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.27.21)
[1541873053] HOST ALERT: dek-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.56.22)
[1541873081] SERVICE ALERT: sadov-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541873085] HOST ALERT: sadov-k2;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1541873094] HOST ALERT: blago-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.71.21)
[1541873094] HOST ALERT: suvor42-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541873119] HOST ALERT: garkav-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541873120] HOST ALERT: stach88-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541873139] SERVICE ALERT: lunach-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541873142] HOST ALERT: lunach-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541873162] HOST ALERT: gavan-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541873183] SERVICE ALERT: stach88-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.54.21)
[1541873218] HOST ALERT: suvor42-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.59.21)
[1541873232] SERVICE ALERT: sadov-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541873243] HOST ALERT: stach88-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.54.21)
[1541873243] SERVICE ALERT: suvor42-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.59.21)
[1541873262] SERVICE ALERT: lunach-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.41.21)
[1541873265] HOST ALERT: sadov-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541873268] HOST ALERT: lunach-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.41.21)
[1541873288] HOST ALERT: gavan-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.48.21)
[1541873320] SERVICE ALERT: dek-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.56.22)
[1541873353] HOST ALERT: sadov-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541873364] HOST ALERT: prosv48-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541873366] SERVICE ALERT: eliz-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541873369] HOST ALERT: eliz-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541873385] SERVICE ALERT: garkav-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541873400] HOST ALERT: eliz-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541873401] HOST ALERT: bkuna-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541873420] HOST ALERT: ushinsk-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541873426] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541873427] SERVICE ALERT: gavan-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.48.21)
[1541873429] HOST ALERT: gasheka-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541873445] SERVICE ALERT: ushinsk-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.25.21)
[1541873489] SERVICE ALERT: eliz-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.22)
[1541873490] HOST ALERT: prosv48-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.67.21)
[1541873492] HOST ALERT: eliz-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.22)
[1541873526] HOST ALERT: eliz-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.11.21)
[1541873527] HOST ALERT: bkuna-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.31.21)
[1541873533] HOST ALERT: sadov-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541873543] HOST ALERT: ushinsk-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.25.21)
[1541873549] SERVICE ALERT: gasheka-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.35.21)
[1541873555] HOST ALERT: gasheka-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.35.21)
[1541873650] SERVICE ALERT: sadov-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541873696] SERVICE ALERT: eliz-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.11.21)
[1541873705] SERVICE ALERT: bkuna-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.31.21)
[1541873807] SERVICE ALERT: prosv48-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.67.21)
[1541873970] SERVICE ALERT: lens22-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541873973] HOST ALERT: lens22-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541873978] HOST ALERT: varsh-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541874093] SERVICE ALERT: lens22-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.61.21)
[1541874096] HOST ALERT: lens22-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.61.21)
[1541874104] HOST ALERT: varsh-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.70.21)
[1541874346] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541874349] HOST ALERT: nevsk104-k4;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541874350] SERVICE ALERT: varsh-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.70.21)
[1541874469] SERVICE ALERT: nevsk104-k4;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.24)
[1541874472] HOST ALERT: nevsk104-k4;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.24)
[1541874594] HOST ALERT: msk197-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541874599] HOST ALERT: nev-k11;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541874653] HOST ALERT: engels51-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541874678] HOST ALERT: kamen-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541874720] HOST ALERT: msk197-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.22)
[1541874751] SERVICE ALERT: msk197-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.22)
[1541874778] HOST ALERT: engels51-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.63.21)
[1541874779] HOST ALERT: nev-k11;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541874804] HOST ALERT: kamen-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.22)
[1541874926] SERVICE ALERT: engels51-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.63.21)
[1541875000] SERVICE ALERT: kamen-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.22)
[1541875021] SERVICE ALERT: nev-k11;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541875152] HOST ALERT: metal-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541875167] HOST ALERT: engels-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541875171] HOST ALERT: tur-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541875278] HOST ALERT: metal-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.69.21)
[1541875293] HOST ALERT: engels-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.40.21)
[1541875297] HOST ALERT: tur-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.57.21)
[1541875448] SERVICE ALERT: metal-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.69.21)
[1541875460] SERVICE ALERT: engels-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.40.21)
[1541875493] SERVICE ALERT: tur-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.57.21)
[1541875805] HOST ALERT: bryanc20-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541875930] HOST ALERT: bryanc20-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.68.21)
[1541876215] SERVICE ALERT: bryanc20-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.68.21)
[1541876394] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541876397] HOST ALERT: prosv87-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541876517] SERVICE ALERT: prosv87-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.21)
[1541876523] HOST ALERT: prosv87-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.21)
[1541876891] HOST ALERT: prosv87-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541876952] HOST ALERT: sad29-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541877017] HOST ALERT: prosv87-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.47.22)
[1541877078] HOST ALERT: sad29-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.43.22)
[1541877191] SERVICE ALERT: prosv87-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.47.22)
[1541877223] SERVICE ALERT: sad29-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.43.22)
[1541877243] SERVICE ALERT: 8line-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541877246] HOST ALERT: 8line-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541877302] HOST ALERT: nev-k2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541877366] SERVICE ALERT: 8line-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.58.22)
[1541877372] HOST ALERT: 8line-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.58.22)
[1541877482] HOST ALERT: nev-k2;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541877612] SERVICE ALERT: nev-k2;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541878764] HOST ALERT: mikro72;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541878804] HOST ALERT: nev-k4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541878828] HOST ALERT: mikro72;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.96 ms
[1541878889] HOST ALERT: nevsk104-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541878984] HOST ALERT: nev-k4;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541879015] HOST ALERT: nevsk104-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.23)
[1541879056] SERVICE ALERT: nev-k4;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541879063] HOST ALERT: nevsky78-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541879131] SERVICE ALERT: nevsk104-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.23)
[1541879148] HOST ALERT: kamen-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541879149] SERVICE ALERT: nevsky78-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.21)
[1541879189] HOST ALERT: nevsky78-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.21)
[1541879274] HOST ALERT: kamen-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.21)
[1541879306] SERVICE ALERT: kamen-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.21)
[1541879486] HOST ALERT: nev-k5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541879487] SERVICE ALERT: nev-k5;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541879666] HOST ALERT: nev-k5;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541879710] HOST ALERT: nev-k6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541879775] SERVICE ALERT: nev-k6;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541879890] HOST ALERT: nev-k6;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541880175] HOST ALERT: msk197-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541880237] HOST ALERT: nevsky78-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541880301] HOST ALERT: msk197-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.46.21)
[1541880303] HOST ALERT: nev-k3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541880363] HOST ALERT: nevsky78-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.28.23)
[1541880386] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541880389] HOST ALERT: nevsky130-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541880393] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541880396] HOST ALERT: nevsk104-k1;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541880483] HOST ALERT: nev-k3;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541880500] SERVICE ALERT: nevsky78-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.28.23)
[1541880509] SERVICE ALERT: nevsky130-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.49.22)
[1541880515] HOST ALERT: nevsky130-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.49.22)
[1541880516] SERVICE ALERT: nevsk104-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.21)
[1541880522] HOST ALERT: nevsk104-k1;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.21)
[1541880526] SERVICE ALERT: kamen-k3;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541880529] HOST ALERT: kamen-k3;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541880587] SERVICE ALERT: msk197-k1;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.46.21)
[1541880603] SERVICE ALERT: nev-k3;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541880649] SERVICE ALERT: kamen-k3;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.44.23)
[1541880655] HOST ALERT: kamen-k3;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.44.23)
[1541880749] SERVICE ALERT: nev-k7;Ping;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1541880779] HOST ALERT: nev-k7;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541880900] SERVICE ALERT: nev-k7;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541880916] HOST ALERT: nev-k9;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541880959] HOST ALERT: nev-k7;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541881096] HOST ALERT: nev-k9;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541881194] HOST ALERT: nev-k1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1541881218] SERVICE ALERT: nev-k9;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541881374] HOST ALERT: nev-k1;DOWN;HARD;3;(Host check timed out after 30.01 seconds)
[1541881499] SERVICE ALERT: nev-k1;Ping;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1541887401] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541887404] HOST ALERT: nevsk104-k2;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541887524] SERVICE ALERT: nevsk104-k2;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.37.22)
[1541887527] HOST ALERT: nevsk104-k2;DOWN;HARD;3;CRITICAL - Host Unreachable (192.168.37.22)